```python
# Initialize Otter
import otter
grader = otter.Notebook("pa5.ipynb")
```

# Programming Assignment 5 Data Analysis and Visualization with Pandas and Matplotlib

In PA4, we explored data cleaning and analysis on the smartphones dataset using the `json` module and basic dictionary operations. This time, we'll leverage the `Pandas` module to work with the same dataset, showcasing how it streamlines the process. Additionally, we’ll create visualizations to bring the data to life.

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import json
```

## Task 1. Load data

1. Load the `data.json` file using the json module.
2. Simplify the data by removing the `"pricing"` field from all entries.
3. Convert the JSON-formatted data into a Pandas DataFrame using the [`json_normalize`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.json_normalize.html) function, specifying `max_level=2`.
4. Save the DataFrame into `df`. 
5. Save the columns of the DataFrame into `df_columns`.

```python
...
df = ...
df_columns = ...
df_columns
```

```python
grader.check("T1")
```

## Task 2: Count the number of phone models for each phone brand

1. Count the number of phone models for each phone brand.
2. Save the results as a Pandas Series in `phone_brand_count`.

Expected results: 

<img src="images/task2.png" width="200">

```python
phone_brand_count = ...
phone_brand_count
```

```python
grader.check("T2")
```

## Task 3. Process the `price` column

`process_price` below is the function from PA4 to process individual raw price. 
1. Use `process_price` to convert the `price` string column into five numerical columns named `price.EUR`, `price.USD`, `price.GBP`, `price.INR`, `price.CAD`, and `price.AUD`. Hint: `apply`, [`rename`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.rename.html), and `json_normalize` will be useful. Make sure to modify `process_price` to handle missing values. 
2. Add these six columns to `df`. Hint: Use `pd.concat` and you may need to specify the `axis` argument.

Expected results: 

<img src="images/task3.png" width="400">

```python
CURRENCY_MAP = {'EUR': 'EUR',
 'INR': 'INR',
 '$': 'USD',
 'A$': 'AUD',
 'C$': 'CAD',
 '£': 'GBP',
 '€': 'EUR',
 '₹': 'INR'
}

def process_price(raw_price_string):
    price_dict = {}

    ...
    
    if "About" in raw_price_string:
        parts = raw_price_string.split()
        amount = float(parts[1])
        currency = parts[2]
        iso_currency = CURRENCY_MAP[currency]
        price_dict[iso_currency] = amount
    else:
        segments = raw_price_string.split(" / ")
        for segment in segments:
            parts = segment.split('\u2009')
            amount = float(parts[1].replace(",", ""))
            currency = parts[0]
            iso_currency = CURRENCY_MAP[currency]
            price_dict[iso_currency] = amount
    return price_dict

...
df[['price.EUR', 'price.USD', 'price.GBP', 'price.INR', 'price.CAD', 'price.AUD']]
```

```python
grader.check("T3")
```

## Task 4. How many missing values are there in the six prices columns? 
Expects a Pandas Series: 
```
price.EUR    1891
price.USD    8301
price.GBP    8357
price.INR    8429
price.CAD    8711
price.AUD    8789
dtype: int64
```

```python
prices_missing_count = ...
prices_missing_count
```

```python
grader.check("T4")
```

## Task 5. Create a histogram of prices under 1500 EUR

1. Create a histogram of prices under 1500 EUR with 30 bins and x, y labels.
2. Save the plot to `price_hist`.

Expected results: 

<img src="images/task5.png" width="400">

```python
price_hist = ...
```

```python
grader.check("T5")
```

## Task 6. Extract release year

If you run `value_counts` on the `'specs.Launch.Status'` column, you'll notice three statuses:
1. Discontinued
2. Cancelled
3. Available. Released year, month

Your tasks: 
1. Extract the release year for each model and add it as a numerical column to `df`, named `release_year`. For models without a status or those marked as Discontinued or Cancelled, leave the value as `NaN`.
2. Save the count of phone models by release year (sorted by year) to `release_year_count`.

Expected results: 
```
release_year
2015.0      1
2016.0     18
2017.0    249
2018.0    359
2019.0    432
2020.0    580
2021.0    659
2022.0    663
2023.0    614
2024.0    517
Name: count, dtype: int64
```

```python
...
release_year_count = ...
release_year_count
```

```python
grader.check("T6")
```

## Task 7. Create a boxplot of prices in EUR by release year

1. Filter the data to include only rows with non-null `release_year` and `price.EUR` values.
2. Group the prices in EUR by release year.
3. Draw a **horizontal** boxplot in `price_boxplot_axes` and save the plot in `price_boxplot`.
4. Annotate the plot with the most expensive model for each year.
5. Label the axes.

Expected results: 

<img src="images/task7.png" width="600">

```python
price_boxplot_axes = plt.subplot()
...
price_boxplot = ...
plt.show()
```

```python
grader.check("T7")
```

## Task 8. Extract network speed options for Apple and Samsung phones

For all Apple and Samsung phones, use the `'specs.Network.Speed'` column (a string of network speed options separated by commas) to create two sets: one listing all network speed options for Apple phones and another for Samsung phones.

Expected results for Apple: 

<img src="images/8apple.png" width="250">

Partial expected results for Samsung: 

<img src="images/8samsung.png" width="600">

```python
apple_network_speed = ...
...
apple_network_speed
```

```python
samsung_network_speed = ...
...
samsung_network_speed
```

```python
grader.check("T8")
```

## Task 9 Count phone models by network speed and release year

From Task 8, we identified four major types of network speed options:
* EV-DO: 3G
* HSPA: 3G
* LTE: 4G
* 5G

Your task is to create two DataFrames, one for Apple phones (`apple_network_df`) and one for Samsung phones (`samsung_network_df`), with the following structure:
* Row indices: release_year (the release year of the phone models).
* Columns: The four network speed options (EV-DO, HSPA, LTE, 5G).
* Cell values: The count of phone models for the given release year and network speed option.

Note: If a phone’s network speed has a specific keyword (e.g., 'EV-DO') as a substring, it falls under that category. Each phone can belong to multiple categories if applicable. Skip any network speed options that do not fall under the four categories.

Expected results for Apple: 

<img src="images/9apple.png" width="230">

Expected results for Samsung: 

<img src="images/9samsung.png" width="230">

```python
...
apple_network_df
```

```python
...
samsung_network_df
```

```python
grader.check("T9")
```

## Task 10. Create heatmaps for network speed options by release year for Apple and Samsung

1. Install Seaborn: Before starting, ensure you have Seaborn installed by running `pip install seaborn`. Activate your virtual environment before running the command. Seaborn, similar to Pandas, uses Matplotlib internally. However, it offers additional features, such as better integration with data and advanced plotting capabilities like heatmaps, which are not natively supported by Matplotlib.

2. Create the Heatmaps:
    * Use the DataFrames you created in the previous task for Apple and Samsung.
    * Set up a subplot with 1 row and 2 columns that share the x-axis and y-axis.
    * The first subplot will be a heatmap for Apple, where:
        * The x-axis represents the four network speed options (EV-DO, HSPA, LTE, 5G).
        * The y-axis represents the release years.
    * The second subplot will be a heatmap for Samsung, using the same axes configuration.
    * Instead of assigning a plot to `ax[?]` directly, use the `ax` parameter in seaborn.heatmap(). For example: `sns.heatmap(data, ax=ax[0])`
    * You can add `cmap='Reds', linewidths=0.5, linecolor='black'` as parameters of the heatmap to distinguish the grid better. 
    * Ensure the subplots have titles (Apple and Samsung) to distinguish them.

Expected results: 

<img src="images/task10.png" width="400">

```python
import seaborn as sns
fig, axes = ...
...
for ax in axes:
    ax.tick_params(axis='y', labelrotation=0)
fig.supxlabel('Network speed')
fig.supylabel('Release year')
plt.show()
```

```python
grader.check("T10")
```

## Conclusion

Through this project, you’ve practiced essential skills in data analysis and visualization. By working with real-world smartphone data, you’ve learned how to process, clean, and extract insights while creating meaningful visualizations. These skills are not limited to this dataset—they can be applied to a wide range of datasets across various domains, enabling you to analyze trends, identify patterns, and effectively communicate findings in any data-driven field.

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for T1 (points: 1)
```
>>> assert df.shape == (8791, 84)
>>> assert df_columns.equals(pd.Index(['phone_brand', 'phone_model', 'price', 'specs.Network.2G bands', 'specs.Network.3G bands', 'specs.Network.4G bands', 'specs.Network.Speed', 'specs.Network.GPRS', 'specs.Network.EDGE', 'specs.Launch.Announced', 'specs.Launch.Status', 'specs.Body.Dimensions', 'specs.Body.Weight', 'specs.Body.Build', 'specs.Body.SIM', 'specs.Display.Type', 'specs.Display.Size', 'specs.Display.Resolution', 'specs.Platform.OS', 'specs.Memory.Card slot', 'specs.Memory.Internal', 'specs.Sound.Loudspeaker', 'specs.Sound.3.5mm jack', 'specs.Comms.WLAN', 'specs.Comms.Bluetooth', 'specs.Comms.Positioning', 'specs.Comms.NFC', 'specs.Comms.Radio', 'specs.Comms.USB', 'specs.Features.Sensors', 'specs.Battery.Type', 'specs.Misc.Colors', 'specs.Platform.Chipset', 'specs.Platform.CPU', 'specs.Platform.GPU', 'specs.Main Camera.Triple', 'specs.Main Camera.Features', 'specs.Main Camera.Video', 'specs.Selfie camera.Single', 'specs.Selfie camera.Video', 'specs.Battery.Charging', 'specs.Misc.Price', 'specs.Display.Protection', 'specs.Battery.Stand-by', 'specs.Battery.Talk time', 'specs.Misc.Models', 'specs.Tests.Loudspeaker', 'specs.Tests.Performance', 'specs.Tests.Display', 'specs.Tests.Battery (old)', 'specs.Main Camera.Single', 'specs.Main Camera.Dual', 'specs.Main Camera.Quad', 'specs.Network.5G bands', 'specs.Network.Technology', 'specs.Selfie Camera.Single', 'specs.Selfie Camera.Video', 'specs.Comms.Infrared port', 'specs.Selfie camera.Features', 'specs.Main Camera.Penta', 'specs.Selfie camera.Dual', 'specs.Tests.Camera', 'specs.Misc.SAR', 'specs.Selfie Camera.Features', 'specs.Features.Languages', 'specs.Body.Keyboard', 'specs.Memory.Phonebook', 'specs.Memory.Call records', 'specs.Features.Messaging', 'specs.Features.Browser', 'specs.Features.Games', 'specs.Features.Java', 'specs.Misc.SAR EU', 'specs.Tests.Audio quality', 'specs.Tests.Battery (new)', 'specs.Battery.Music play', 'specs.Sound.Alert types', 'specs.Features.Clock', 'specs.Features.Alarm', 'specs.Selfie camera.No', 'specs.Selfie camera.Triple', 'specs.Main Camera.Dual or Triple', 'specs.Main Camera.Five', 'specs.Selfie Camera.Dual'], dtype='object'))
```


Tests for T10 (points: 1)
```
>>> ax1, ax2 = (axes[0], axes[1])
>>> assert ax1.get_title() == 'Apple'
>>> assert ax2.get_title() == 'Samsung'
>>> assert np.array_equal(ax1.get_xticks(), np.array([0.5, 1.5, 2.5, 3.5]))
>>> assert np.array_equal(ax2.get_xticks(), np.array([0.5, 1.5, 2.5, 3.5]))
>>> assert np.array_equal(ax1.get_yticks(), np.array([0.5, 1.5, 2.5, 3.5, 4.5, 5.5, 6.5, 7.5, 8.5, 9.5]))
>>> assert np.array_equal(ax2.get_yticks(), np.array([0.5, 1.5, 2.5, 3.5, 4.5, 5.5, 6.5, 7.5, 8.5, 9.5]))
>>> vmin1, vmax1 = (ax1.collections[0].colorbar.vmin, ax1.collections[0].colorbar.vmax)
>>> vmin2, vmax2 = (ax2.collections[0].colorbar.vmin, ax2.collections[0].colorbar.vmax)
>>> assert vmin1 == np.float64(0.0)
>>> assert vmax1 == np.float64(13.0)
>>> assert vmin1 == np.float64(0.0)
>>> assert vmax2 == np.float64(55.0)
```


Tests for T2 (points: 1)
```
>>> phone_brand_count_index = list(phone_brand_count.index)
>>> assert phone_brand_count_index[:10] == ['samsung', 'lg', 'motorola', 'nokia', 'vivo', 'huawei', 'xiaomi', 'alcatel', 'zte', 'oppo']
>>> assert phone_brand_count_index[12:24] == ['lenovo', 'honor', 'realme', 'asus', 'sony', 'tecno', 'infinix', 'apple', 'ulefone_', 'doogee', 'blackview', 'meizu']
>>> assert phone_brand_count_index[29:] == ['umidigi', 'oukitel', 'itel', 'google', 'microsoft', 'cat', 'nothing']
>>> assert phone_brand_count.values == pd.array([1400, 667, 634, 585, 476, 460, 412, 409, 387, 318, 289, 289, 249, 227, 221, 200, 160, 153, 140, 128, 99, 92, 80, 78, 77, 77, 77, 75, 75, 71, 55, 39, 33, 32, 22, 5])
```


Tests for T3 (points: 1)
```
>>> expected_prices = pd.read_csv('prices.csv')
>>> assert df[['price.EUR', 'price.USD', 'price.GBP', 'price.INR', 'price.CAD', 'price.AUD']].equals(expected_prices)
```


Tests for T4 (points: 1)
```
>>> assert prices_missing_count.index.equals(pd.Index(['price.EUR', 'price.USD', 'price.GBP', 'price.INR', 'price.CAD', 'price.AUD']))
>>> assert prices_missing_count.values == pd.array([1891, 8301, 8357, 8429, 8711, 8789])
```


Tests for T5 (points: 1)
```
>>> assert price_hist.get_xlabel() == 'Prices (EUR)'
>>> assert price_hist.get_ylabel() == 'Count'
>>> assert [p.get_height() for p in price_hist.patches] == [np.float64(458.0), np.float64(1278.0), np.float64(1329.0), np.float64(1049.0), np.float64(767.0), np.float64(595.0), np.float64(333.0), np.float64(256.0), np.float64(181.0), np.float64(91.0), np.float64(105.0), np.float64(73.0), np.float64(65.0), np.float64(49.0), np.float64(45.0), np.float64(25.0), np.float64(33.0), np.float64(24.0), np.float64(25.0), np.float64(10.0), np.float64(20.0), np.float64(10.0), np.float64(7.0), np.float64(8.0), np.float64(8.0), np.float64(4.0), np.float64(9.0), np.float64(1.0), np.float64(5.0), np.float64(7.0)]
```


Tests for T6 (points: 1)
```
>>> assert release_year_count.index.equals(pd.Index([2015.0, 2016.0, 2017.0, 2018.0, 2019.0, 2020.0, 2021.0, 2022.0, 2023.0, 2024.0], dtype='float64', name='release_year'))
>>> assert release_year_count.values == pd.array([1, 18, 249, 359, 432, 580, 659, 663, 614, 517])
```


Tests for T7 (points: 1)
```
>>> assert price_boxplot_axes.get_xlabel() == 'Prices (EUR)'
>>> assert price_boxplot_axes.get_ylabel() == 'Release year'
>>> assert [t.get_text() for t in price_boxplot_axes.texts] == ['Samsung Galaxy Tab A 8.0 & S Pen (2015)', 'Lenovo Yoga Tab 3 Pro', 'Apple Watch Edition Series 3', 'Oppo Find X Lamborghini', 'Huawei Mate X', 'Huawei Mate Xs', 'Huawei Mate X2', 'Honor Magic Vs', 'Huawei Watch Ultimate Design', 'Honor Magic V2 RSR Porsche Design']
>>> student_medians = [line.get_xdata() for line in price_boxplot['medians']]
>>> expected_medians = [np.array([220.0, 220.0]), np.array([140.0, 140.0]), np.array([200.0, 200.0]), np.array([220.0, 220.0]), np.array([230.0, 230.0]), np.array([200.0, 200.0]), np.array([200.0, 200.0]), np.array([230.0, 230.0]), np.array([235.995, 235.995]), np.array([259.995, 259.995])]
>>> for i in range(len(student_medians)):
...     assert np.array_equal(student_medians[i], expected_medians[i])
>>> student_whiskers = [line.get_xdata() for line in price_boxplot['whiskers']]
>>> expected_whiskers = [np.array([220.0, 220.0]), np.array([220.0, 220.0]), np.array([100.0, 70.0]), np.array([190.0, 220.0]), np.array([150.0, 10.0]), np.array([300.0, 520.0]), np.array([150.0, 10.0]), np.array([350.0, 650.0]), np.array([150.0, 10.0]), np.array([400.0, 760.0]), np.array([130.0, 20.0]), np.array([356.2425, 690.0]), np.array([130.0, 10.0]), np.array([325.0, 610.0]), np.array([150.0, 10.0]), np.array([370.0, 700.0]), np.array([150.0, 20.0]), np.array([409.9975, 799.99]), np.array([160.0, 20.0]), np.array([459.985, 900.0])]
>>> for i in range(len(student_whiskers)):
...     assert np.array_equal(student_whiskers[i], expected_whiskers[i])
>>> student_caps = [line.get_xdata() for line in price_boxplot['caps']]
>>> expected_caps = [np.array([220.0, 220.0]), np.array([220.0, 220.0]), np.array([70.0, 70.0]), np.array([220.0, 220.0]), np.array([10.0, 10.0]), np.array([520.0, 520.0]), np.array([10.0, 10.0]), np.array([650.0, 650.0]), np.array([10.0, 10.0]), np.array([760.0, 760.0]), np.array([20.0, 20.0]), np.array([690.0, 690.0]), np.array([10.0, 10.0]), np.array([610.0, 610.0]), np.array([10.0, 10.0]), np.array([700.0, 700.0]), np.array([20.0, 20.0]), np.array([799.99, 799.99]), np.array([20.0, 20.0]), np.array([900.0, 900.0])]
>>> for i in range(len(student_caps)):
...     assert np.array_equal(student_caps[i], expected_caps[i])
```


Tests for T8 (points: 1)
```
>>> assert apple_network_speed == {'5G', 'EV-DO Rev.A 3.1 Mbps', 'HSPA', 'HSPA 14.4/2 Mbps', 'HSPA 14.4/5.76 Mbps', 'HSPA 42.2/5.76 Mbps', 'HSPA 7.2/0.384 Mbps', 'HSPA 7.2/5.76 Mbps', 'LTE', 'LTE (2CA) Cat6 300/50 Mbps', 'LTE (3CA) Cat12 600/150 Mbps', 'LTE (3CA) Cat9 450/50 Mbps', 'LTE (4CA) Cat16 1024/150 Mbps', 'LTE (5CA) Cat16 1024/150 Mbps', 'LTE (CA)', 'LTE (CA) Cat16 1024/150 Mbps', 'LTE 1.6 Gbps DL', 'LTE Cat3 100/50 Mbps', 'LTE Cat4 150/50 Mbps', 'LTE Cat6 300/50 Mbps'}
>>> assert samsung_network_speed == {'153.2 Kbps', '384 kbps', '42.2(21.1)/5.76 Mbps (region dependent)', '5G', '5G (2+ Gbps DL)', '5G (5+ Gbps DL)', 'CDMA2000 1xEV-DO', 'EV-DO Rev. 0', 'EV-DO Rev. A 3.1 Mbps', 'EV-DO Rev.0', 'EV-DO Rev.A 3.1 Mbps', 'HSPA', 'HSPA 1.8 Mbps', 'HSPA 14.4/5.76 Mbps', 'HSPA 21.1/3.6 Mbps', 'HSPA 21.1/5.76 Mbps', 'HSPA 21.1/5.76 Mbps (I9200)', 'HSPA 3.6/0.384 Mbps', 'HSPA 3.6/2 Mbps', 'HSPA 42.2(14.4 - 3G)/5.76 Mbps', 'HSPA 42.2(LTE)/21.1(3G)/5.76 Mbps', 'HSPA 42.2/11.5 Mbps', 'HSPA 42.2/21/5.76 Mbps (region dependent)/ LTE Cat3 100/50 Mbps (SM-G7105)', 'HSPA 42.2/5.76 Mbps', 'HSPA 42.2/5.76 Mbps (T321)/ LTE Cat4 150/50 Mbps (T325)', 'HSPA 7.2/0.384 Mbps', 'HSPA 7.2/0.384 Mbps (GT-C3595)', 'HSPA 7.2/2 Mbps', 'HSPA 7.2/5.76 Mbps', 'LTE', 'LTE (2CA)', 'LTE (2CA) Cat12 600/150 Mbps', 'LTE (2CA) Cat12 600/50 Mbps', 'LTE (2CA) Cat13 400/50 Mbps', 'LTE (2CA) Cat6 300/50 Mbps', 'LTE (2CA) Cat6 300/50 Mbps/ LTE (3CA) Cat9 450/50 Mbps', 'LTE (2CA) Cat6 400/50 Mbps', 'LTE (3CA)', 'LTE (3CA) Cat11 600/50 Mbps', 'LTE (3CA) Cat11 600/75 Mbps', 'LTE (3CA) Cat12 600/150 Mbps', 'LTE (3CA) Cat6 300/50 Mbps', 'LTE (3CA) Cat6 300/50 Mbps or Cat4 150/50 Mbps', 'LTE (3CA) Cat9 450/50 Mbps', 'LTE (3CA) Cat9/Cat10 or LTE (3CA) Cat12 600/150 Mbps', 'LTE (4CA) Cat12 600/150 Mbps', 'LTE (5CA) Cat16 1000/150 Mbps', 'LTE (5CA) Cat16 1024/150 Mbps', 'LTE (5CA) Cat18 1200/200 Mbps', 'LTE (5CA) Cat19 1600/200 Mbps', 'LTE (5CA) Cat19 1800/200 Mbps', 'LTE (6CA)', 'LTE (6CA) Cat18 1200/150 Mbps', 'LTE (6CA) Cat18 1200/200 Mbps', 'LTE (6CA) Cat20 2000/150 Mbps', 'LTE (6CA) Cat20 2000/200 Mbps', 'LTE (7CA) Cat20 2000/150 Mbps', 'LTE (7CA) Cat20 2000/200 Mbps', 'LTE (CA)', 'LTE (CA) Cat18 1200/150 Mbps', 'LTE (T999L)', 'LTE (optional)', 'LTE (up to 7CA)', 'LTE Cat16 1024/75 Mbps', 'LTE Cat3 100/50 Mbps', 'LTE Cat3 100/50 Mbps (I9205)', 'LTE Cat3 100/50 Mbps (SM-T215)', 'LTE Cat4 150/50 Mbps', 'LTE Cat4 150/50 Mbps (Exynos 5433)', 'LTE Cat4 150/50 Mbps or LTE (2CA) Cat6 300/50 Mbps', 'LTE Cat6 300/50 Mbps', 'LTE Cat6 300/50 Mbps/ LTE (3CA) Cat9 450/50 Mbps', 'LTE-А', 'No', 'Yes'}
```


Tests for T9 (points: 1)
```
>>> assert np.array_equal(apple_network_df.index, pd.RangeIndex(start=2015, stop=2025, step=1))
>>> assert np.array_equal(samsung_network_df.index, pd.RangeIndex(start=2015, stop=2025, step=1))
>>> assert apple_network_df.columns.equals(pd.Index(['EV-DO', 'HSPA', 'LTE', '5G'], dtype='object'))
>>> assert samsung_network_df.columns.equals(pd.Index(['EV-DO', 'HSPA', 'LTE', '5G'], dtype='object'))
>>> assert np.array_equal(apple_network_df.to_numpy(), np.array([[0, 0, 0, 0], [0, 0, 0, 0], [5, 8, 8, 0], [6, 8, 8, 0], [4, 9, 9, 0], [7, 13, 13, 4], [6, 11, 11, 7], [5, 13, 13, 9], [4, 7, 7, 4], [4, 10, 10, 8]]))
>>> assert np.array_equal(samsung_network_df.to_numpy(), np.array([[0, 1, 1, 0], [0, 0, 0, 0], [0, 19, 19, 0], [0, 30, 30, 0], [0, 43, 43, 5], [0, 55, 55, 18], [0, 43, 43, 17], [0, 37, 37, 21], [0, 32, 32, 23], [0, 28, 28, 19]]))
```
