```python
# Initialize Otter
import otter
grader = otter.Notebook("pa1.ipynb")
```

```python
# change the following if needed
# partner: none
```

# Programming Assignment 1

## How to complete this assignment? 
1. The first cell of this notebook is created by the otter grader. Please run this cell. 
```python
# Initialize Otter
import otter
grader = otter.Notebook("pa1.ipynb")
```

2. Complete the missing parts, indicated by `...`, in each question.


3. You can check each question using the `grader` object by running:
```python
grader.check("q1")
```

4. You can check all questions at once using the `grader` object by running:
```python
grader.check_all()
```


## Section 1. Installation and Virtual Environment 

**Note:** 
Before you start, follow the installation and virtual environment tutorials from [Lecture 1](https://github.com/yiyins2/cs368-python-fa24/tree/main/01-installation-version-control). Make sure this notebook is run within the virtual environment.

### Question 1. Display the Python version using `!python --version`.

In Jupyter Notebooks, prefixing a shell command with `!` allows it to execute and return results within the notebook just as it would in the shell.

**Task:** Assign the Python version into the `python_version` variable using `!python --version`. 

**Expected Output:** Since the notebook runs within the virtual environment you created for the course, which is based on Python 3.12.x, it should output ['Python 3.12.x'].

```python
python_version = ...
python_version
```

```python
grader.check("Q1")
```

### Question 2. Display the path of the excutable using `sys.executable`.

In Python, you can import packages using the syntax `import <package_name>`. Begin by importing the sys library with `import sys`. This library, included in Python's standard library, provides functions and variables that directly interact with the Python interpreter.

One attribute of the sys package is `sys.executable`, which indicates the path to the Python executable that is running the current script. Attributes and functions of a package can be accessed using the dot notation in the format `<package_name>.<attribute_or_function>`. 

**Tasks:**
1. Import the `sys` library.
2. Assign the path to the Python executable to the `executable` variable using the `sys.executable` attribute.

**Expected Output:** Again, since the notebook runs within the virtual environment you created, it should output the path of the executable within the local copy of CPython.

```python
...
executable = ...
executable
```

```python
grader.check("Q2")
```

### Question 3. Install packages using `requirements.txt` and check Matplotlib version. 

In Python, to simplify references to a package throughout your code, you can import it using a shorthand alias with the syntax `import <package_name> as <alias>`. Once imported with alias, you should refer to the package using the alias. To access this package's attributes or functions, use `<alias>.<attribute_or_function>`. 

**Tasks:**

1. Inside the folder for `PA1`, you'll find a requirements.txt file that lists some of the packages we'll use in this course, along with their required versions. Install these packages using the command `pip install -r requirements.txt`.
2. Import `matplotlib` using the alias `plt` and then use this alias to display matplotlib's version. You may refer to [this documentation page](https://matplotlib.org/stable/install/index.html#obtaining-matplotlib-version) to locate the appropriate attribute for obtaining the version.

**Expected Output:** The output should be '3.9.2', which is the latest stable version of Pandas.

```python
...
matplotlib_version = ...
matplotlib_version
```

```python
grader.check("Q3")
```

### Question 4. Import a self-defined module. 

Apart from executing a `.py` script using the commend line, you can also import and use it in Jupyter Notebook. 

**Tasks:**
1. Create a `.py` file named `pa1` in the same directory as this notebook. Then, write the following lines in `pa1.py`:

```python
print('hello, world!')

def info(): 
   return 'CS368 PA1'
```

Note that Python functions are defined in the following format:

```python
def function_name(parameters):
    return ...
```
   More details of functions will be covered in later lectures.

2. A single `.py` file is often referred to as a module. Similar to how you import libraries, you can import a `module.py` using `import module`. Import the `pa1.py` you just created.

3. Invoke the `info` function defined in `pa1.py`. 

**Expected Output:** You should see 'hello, world!' printed and then 'CS368 PA1' displayed in the output cell.

```python
...
info_return = ...
info_return 
```

```python
grader.check("Q4")
```

### Question 5. Auto-reload the module.

**Tasks:**
1. Add another print statement to your `info()` function before the return statement.

2. Invoke the `info` function again. You won't see the new print statement because new changes made in the module won't automatically reflect in the notebook. You can choose to restart the kernel and import the module again whenever you made changes, or ...

You can add the following two lines at the beginning of your notebook and restart the kernel once:
```python
%load_ext autoreload
%autoreload 2
```

These two lines ensure that all module changes are automatically reloaded before any code execution, reflecting any updates immediately without needing to restart the kernel.

```python
info_return2 = ...
info_return2
```

```python
grader.check("Q5")
```

## Section 2. Basic Data Types

### Question 6. Inspect type.
**Task:** Get the type of `q6_var` and store it in `q6_type`. 

```python
q6_var = 'hello, world!'
q6_type = ...
q6_type
```

```python
grader.check("Q6")
```

### Question 7. Method resolution order.
**Task:** Store the method resolution order for the `str` type in `str_mro`.

Note `str`'s parent class is `object`. Everything is Pythin is an object!

```python
str_mro = ...
str_mro
```

```python
grader.check("Q7")
```

### Question 8. Formatting floats.
**Task:** Format the float 3.68 to a string with 30 decimal places.

```python
format30 = ...
format30
```

```python
grader.check("Q8")
```

### Question 9. Float arithmetics and comparison.
**Task:** Given x = 0.1 + 0.2 + 0.3 and y = 0.6, does the values of x and y equal? 

```python
x = 0.1 + 0.2 + 0.3
y = 0.6
x_y_comparison_result = ...
x_y_comparison_result
```

```python
grader.check("Q9")
```

### Question 10. Implement quadratic formula. 
**Task:** Given a quadratic equation $ax^2 + bx + c = 0$, complete the following function that returns the two roots of the equation using the [quadratic formula](https://en.wikipedia.org/wiki/Quadratic_formula). Assume the equation given is solvable. `root1` is the bigger root. 

```python
def quadratic_formula(a, b, c): 
    discriminant = ...
    root1 = ...
    root2 = ...
    return root1, root2
quadratic_formula(1, 4, 3)
```

```python
grader.check("Q10")
```

### Question 11. Implement XOR.
**Task:** Use `not`, `and` and `or` to implement the following [exclusive or](https://en.wikipedia.org/wiki/Exclusive_or) function. 

```python
def xor(operand1, operand2): 
    result = ...
    return result
```

```python
grader.check("Q11")
```

### Question 12. Represent `time` in terms of `hours`, `minutes`, and `seconds`. 

**Task:**
Complete the following `time_converter` function, where the argument passed, `time`, is a non-negative integer measured in seconds. Convert `time` to the form of `x hours y minutes and z seconds`, ensuring that `hours`, `minutes`, and `seconds` are integers. For example, 10,000 seconds is 2 hours 46 minutes and 40 seconds. 

```python
def time_converter(time): 
    hours = ...
    minutes = ...
    seconds = ...
    return hours, minutes, seconds

time_converter(10000)
```

```python
grader.check("Q12")
```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for Q1 (points: 1)
```
>>> assert python_version[0].startswith('Python 3.12')
```


Tests for Q10 (points: 1)
```
>>> import math
>>> r1, r2 = quadratic_formula(1, 4, 3)
>>> assert math.isclose(r1, -1.0) and math.isclose(r2, -3.0)
>>> r1, r2 = quadratic_formula(2, 3, -4)
>>> assert math.isclose(r1, 0.8507810593582121) and math.isclose(r2, -2.350781059358212)
>>> r1, r2 = quadratic_formula(4, 10, 5)
>>> assert math.isclose(r1, -0.6909830056250525) and math.isclose(r2, -1.8090169943749475)
```


Tests for Q11 (points: 1)
```
>>> assert not xor(True, True)
>>> assert not xor(False, False)
>>> assert xor(True, False)
>>> assert xor(False, True)
```


Tests for Q12 (points: 1)
```
>>> assert time_converter(0) == (0, 0, 0)
>>> assert time_converter(10) == (0, 0, 10)
>>> assert time_converter(100) == (0, 1, 40)
>>> assert time_converter(300) == (0, 5, 0)
>>> assert time_converter(5000) == (1, 23, 20)
>>> assert time_converter(10000) == (2, 46, 40)
```


Tests for Q2 (points: 1)
```
>>> import platform
>>> if platform.system() == 'Darwin':
...     assert executable.endswith('/bin/python3.12') and 'local' not in executable
... elif platform.system() == 'Windows':
...     assert executable.endswith('\\Scripts\\python.exe') and 'AppData' not in executable
... elif platform.system() == 'Linux':
...     assert executable.endswith('/bin/python') and (not executable.startswith('/usr'))
... else:
...     print('Unknown Operating System')
```


Tests for Q3 (points: 1)
```
>>> assert matplotlib_version == '3.9.2'
>>> assert 'plt' in globals()
```


Tests for Q4 (points: 1)
```
>>> assert info_return == 'CS368 PA1'
```


Tests for Q5 (points: 1)
```
>>> assert info_return2 == 'CS368 PA1'
```


Tests for Q6 (points: 1)
```
>>> assert q6_type == str
```


Tests for Q7 (points: 1)
```
>>> assert str_mro == [str, object]
```


Tests for Q8 (points: 1)
```
>>> assert format30 == '3.680000000000000159872115546023'
```


Tests for Q9 (points: 1)
```
>>> assert not x_y_comparison_result
```
