```python
# Initialize Otter
import otter
grader = otter.Notebook("pa3.ipynb")
```

```python
import random
```

# Programming Assignment 3 Ecosystem Simulation

In this assignment, you’ll build an ecosystem simulation using Object-Oriented Programming (OOP) principles. The simulation will use inheritance to model plants and animals—herbivores, carnivores, and omnivores. You will define interactions, including special methods, between organisms of the same type and different types, such as growing, hunting, eating, and reproducing.

## Section 1: The `Organism` Base Class
The `Organism` class is a base class representing any living entity in the ecosystem. It defines common attributes and behaviors that all plants and animals will inherit. Each organism has a name, a species and a health attribute, along with methods for interacting with the environment, such as checking if the organism is alive.

#### Attributes:
**Note**: Make sure to adapt the following names into private or protected attributes.
* `name`: The **private** attribute representing the name of the organism.
* `species`: The **private** attribute representing the species of the organism.
* `health`: The **protected** attribute representing the health of the organism, initialized to a default value of 100 if not provided.

#### Methods:

**Note:** `...` underneath are placeholders for parameters, please fill it yourself according to the examples underneath. 

1. `__init__(...)`:
    * Initializes the `Organism` object with a name (required), a species (required) and a health value (optional, default 100).
2. `get_name(...)`:
    * Returns the name of the organism.
3. `get_species(...)`:
    * Returns the species of the organism.
4. `get_health(...)`:
    * Returns the health of the organism.
5. `set_health(...)`:
    * Updates the organism's health to a new health. If the organism is already dead (i.e., health is 0 or less), it raises a `ValueError` with the message `"Cannot update health for {organism_name}, organism is dead."`, where `{organism_name}` is the organism's name.
6. `is_alive(...)`:
    * Returns True if the organism's health is greater than 0; otherwise, returns False.

```python
...
```

```python
# Example
organism = Organism(name="Oak1", species="Oak Tree", health=80)

# Test attributes' initialization
print(organism.get_name())    # Expected: "Oak1"
print(organism.get_species()) # Expected: "Oak Tree"
print(organism.get_health())  # Expected: 80
print(organism.is_alive())    # Expected: True

# Test updating health
organism.set_health(50)
print(organism.get_health())  # Expected: 50
```

```python
grader.check("S1")
```

## Section 2: The `Plant` Class
The `Plant` class represents a plant in the ecosystem that inherits the `Organism` class. It introduces a private `growth_rate` attribute, which determines how much the plant's health increases over time. The class includes methods for accessing the plant's growth rate, as well as simulating plant growth.

#### Attributes:
* Inherits all attributes from the `Organism` class.
* `growth_rate`: The **private** attribute representing the rate at which the plant grows, initialized to a default value of 5 if not provided.

#### Methods:
1. `__init__(...)`:
    * Initializes the `Plant` object with a name (required), a species (required), health (optional, default 100), and growth rate (optional, default 5).
    * **Note**: Remember **NOT** to recreate attributes that are already present in the super class. 
2. `get_growth_rate(...)`:
    * Returns the plant’s growth rate.
4. `grow(...)`:
    * Increases the plant’s health by its growth rate. This simulates the plant growing over time.
    * **Note:** Please use the `set_health(...)` method from the `Organism` class to update the plant’s health.
6. `__str__(...)`:
    * Returns a user-friendly string in the form of `N (species: S, health: H, growth rate: G)`, where `N` is the plant's name, `S` the species, `H` the health, and `G` the growth rate.
7. `__repr__(...)`:
    * Returns a detailed string for debugging in the form of `Plant(name=N, species=S, health=H, growth_rate=G)`.
8. `__eq__(...)`:
    * Compares two plants for equality based on their name, species, health, and growth rate. Returns True if all are equal, otherwise returns False.

```python
...
```

```python
# Example
plant = Plant(name="Oak1", species="Oak Tree", health=80, growth_rate=10)

# Test attributes' initialization
print(plant.get_name())        # Expected: "Oak1"
print(plant.get_species())     # Expected: "Oak Tree"
print(plant.get_health())      # Expected: 80
print(plant.get_growth_rate()) # Expected: 10

# Test grow method
plant.grow()
print(plant.get_health())      # Expected: 90 (80 + 10)

# Test string representation
print(plant)                   # Expected: "Oak1 (species: Oak Tree, health: 90, growth rate: 10)"
print(repr(plant))             # Expected: "Plant(name=Oak1, species=Oak Tree, health=90, growth_rate=10)"

# Test __eq__()
plant2 = Plant(name="Oak2", species="Oak Tree", health=80, growth_rate=5)
print(plant == plant2)        # Expected: False
```

```python
grader.check("S2")
```

## Section 3: The `Animal` Class
The `Animal` class represents animals in the ecosystem and inherits the `Organism` class. It introduces a private `aging_rate` attribute, which determines how much an animal’s health decreases over time. The class includes methods to manage aging and represent the animal’s status.

#### Attributes:
* Inherits all attributes from the `Organism` class.
* `aging_rate`: The **private** attribute representing the rate at which the animal ages, initialized to a default value of 1 if not provided.

#### Methods:
1. `__init__(...)`:
    * Initializes the `Animal` object with a name (required), a species (required), health (optional, default 100), and aging rate (optional, default 1).
    * **Note**: Remember **NOT** to recreate attributes that are already present in the super class. 
2. `get_aging_rate(...)`:
    * Returns the animal’s aging rate.
4. `age(...)`:
    * Decreases the animal’s health by its aging rate. This simulates the animal aging over time.
    * **Note:** Please use the `set_health(...)` method from the `Organism` class to update the animal’s health. 
5. `__str__(...)`:
    * Returns a user-friendly string in the form of `N (species: S, health: H, aging rate: G)`, where `N` is the animal's name, `S` the species, `H` the health, and `G` the aging rate.
7. `__repr__(...)`:
    * Returns a detailed string for debugging in the form of `Animal(name=N, species=S, health=H, aging_rate=G)`.

```python
...
```

```python
# Example
animal = Animal(name="Luna", species="Cat", health=100, aging_rate=5)

# Test attributes' initialization
print(animal.get_name())        # Expected: "Luna"
print(animal.get_species())     # Expected: "Cat"
print(animal.get_health())      # Expected: 100
print(animal.get_aging_rate())  # Expected: 5

# Test age method
animal.age()
print(animal.get_health())      # Expected: 95 (100 - 5)

# Test string representation
print(animal)                   # Expected: "Luna (species: Cat, health: 95, aging rate: 5)"
print(repr(animal))             # Expected: "Animal(name=Luna, species=Cat, health=95, aging_rate=5)
```

```python
grader.check("S3")
```

## Section 4: The `Herbivore` Class
The `Herbivore` class represents animals that eat plants in the ecosystem. It inherits from the `Animal` class and introduces the behavior of eating plants to increase its health.

#### Attributes:
* Inherits all attributes from the `Animal` class, including `name`, `species`, `health`, and `aging_rate`.
* `meal_size`: The **private** attribute representing how much health the herbivore gains and how much health the plant loses per "meal", initialized to a default value of 20.

#### Methods:
1. `__init__(...)`:
    * Initializes a `Herbivore` object with a name, species, health (optional, default 100), aging rate (optional, default 1), and meal size (optional, default 20).
    * **Note**: Remember **NOT** to recreate attributes that are already present in the super class. 
2. `get_meal_size(...)`
    * Returns the current meal size
3. `set_meal_size(...)`:
    * Updates the herbivore’s meal size to a new value.
4. `eat_plant(self, plant)`:
    * If the plant is dead (i.e., health ≤ 0), the method returns, and no action is taken.
    * If the plant is alive:
        * The herbivore gains health equal to the meal size, and the plant loses health by that much.
        * If the plant's health is less than the meal size, the herbivore's health increase by the amount of the plant's remaining health, and set the plant's health to 0.

```python
...
```

```python
# Example
herbivore = Herbivore(name="Deer1", species="Deer", health=100, aging_rate=2, meal_size=10)
plant = Plant(name="Grass1", species="Grass", health=50, growth_rate=5)

# Test attributes' initialization
print(herbivore.get_name())        # Expected: "Deer1"
print(herbivore.get_species())     # Expected: "Deer"
print(herbivore.get_health())      # Expected: 100
print(plant.get_health())          # Expected: 50

# Test eating a plant with sufficient health
herbivore.eat_plant(plant)         
print(herbivore.get_health())      # Expected: 110 (100 + 10)
print(plant.get_health())          # Expected: 40 (50 - 10)

# Test eating the plant with less than meal size health
plant.set_health(5)                # Plant now has less than the meal size
herbivore.eat_plant(plant)      
print(herbivore.get_health())      # Expected: 115 (110 + 5)
print(plant.get_health())          # Expected: 0

# Test attempting to eat a dead plant
herbivore.eat_plant(plant)     
print(herbivore.get_health())      # Expected: 115 (no change)
```

```python
grader.check("S4")
```

## Section 5: The `Carnivore` Class
#### Attributes:
* Inherits all attributes from the Animal class, including name, species, health, and aging_rate.
* `success_rate`: The **private** attribute representing the probability (between 0 and 1) of a successful hunt, initialized to a default value of 0.5.

#### Methods:
1. `__init__(...):`
    * Initializes a `Carnivore` object with a name, species, health (optional, default 100), aging rate (optional, default 1), and success rate (optional, default 0.5).
    * **Note**: Remember **NOT** to recreate attributes that are already present in the super class. 
2. `get_success_rate(...)`:
    * Returns the current success rate of the carnivore.
3. `set_success_rate(...)`:
    * Updates the carnivore’s success rate to a new value.
4. `hunt(self, prey)`:
    * The success of the hunt is determined by the carnivore’s success_rate:
    * If the hunt is successful, set the prey’s health to 0 (i.e., the prey dies), and the carnivore gains all the prey’s health.
        * Hint: Code to determine if a hunt is successful:
          ```python
          import random
          if random.random() < success_rate: # random.random() returns a random value between 0 to 1. 
              # success
          ```
    * If the hunt fails, no health is changed for either the carnivore or the prey.
    * If the prey is already dead, the hunt fails, and no further action is taken.

```python
...
```

```python
# Example
carnivore = Carnivore(name="Lion1", species="Lion", health=100, aging_rate=2, success_rate=1)
prey = Herbivore(name="Zebra1", species="Zebra", health=80, aging_rate=2, meal_size=3)

# Test attributes' initialization
print(carnivore.get_name())        # Expected: "Lion1"
print(carnivore.get_species())     # Expected: "Lion"
print(carnivore.get_health())      # Expected: 100
print(prey.get_health())           # Expected: 80

# Test hunt method
carnivore.hunt(prey)
print(carnivore.get_health())      # Expected: 180 (100 + 80)
print(prey.get_health())           # Expected: 0 (prey is dead)

# Test failed hunt
carnivore_0_success = Carnivore(name="Lion0", species="Lion", health=100, aging_rate=2, success_rate=0)
carnivore_0_success.hunt(prey) 
print(carnivore_0_success.get_health())      # Expected: 100
```

```python
grader.check("S5")
```

## Section 6: The `Omnivore` Class
## The `Omnivore` Class
The `Omnivore` class represents animals that can eat both plants and other animals. It inherits from both the `Herbivore` and `Carnivore` classes, giving it the ability to both hunt other animals and eat plants. 

#### Attributes:
* Inherits all attributes from both the Herbivore and Carnivore classes, including name, species, health, aging_rate, and success_rate for hunting.

#### Methods:
1. `__init__(...)`:
    * Initializes an Omnivore object with a name, species, health (optional, default 100), aging rate (optional, default 1), meal size (optional, default 20), and success rate for hunting (optional, default 0.5).
    * **Note**: Remember **NOT** to recreate attributes that are already present in the super classes. 

```python
...
```

```python
# Set the seed for random to ensure reproducibility
random.seed(50)

omnivore = Omnivore("Bear", "Bear", health=100, aging_rate=2, meal_size=10, success_rate=0.8)
plant = Plant("Berries", "Berry", 50, 5)
prey = Herbivore("SmallFish", "Fish", 40, 3)

# Test attributes' initialization
print(omnivore.get_name())        # Expected: "Bear"
print(omnivore.get_species())     # Expected: "Bear"
print(omnivore.get_health())      # Expected: 100

# Test eating plants
omnivore.eat_plant(plant)         
print(omnivore.get_health())      # Expected: 110 (100 + 10)
print(plant.get_health())         # Expected: 40 (50 - 10)

# Test hunting prey
omnivore.hunt(prey)               
print(omnivore.get_health())      # Expected: 150 (110 + 40)
print(prey.get_health())          # Expected: 0 (prey is dead)
```

```python
grader.check("S6")
```

## Section 7: Add Reproduction Methods
Now, let's add an extra functionality of reproduction in the above classes. You will implement the `reproduce_helper` method in the `Organism` class, which handles the type checks and offspring health calculation. This helper method will be invoked in the `reproduce` methods of both the `Plant` and `Animal`'s classes.

You can go back to the previous sections and append the new methods to the class definitions. 

* Implement `reproduce_helper(self, other)` in `Organism` Class. 
    1. The method checks if the types of `self` and `other` are same. If not, raise a `TypeError` with the message of `"Can only reproduce with another organism of the same type."`
    2. checks if the species of `self` and `other` are same. If not, raise a `TypeError` with the message of `"Can only reproduce with another organism of the same species."`
    3. If both conditions are satisfied, the method returns the offspring’s health calculated as `(self's health + other's health) / 6`. 

* Implement `reproduce(self, other, new_name)` in `Plant`, `Herbivore`, `Carnivore`, and `Omnivore` Class.
    1. Use the `reproduce_helper` method from the `Organism` class to check the types and species, and obtain the offspring’s health.
    2. Returns a new offspring object of its own class with
        * the given `new_name`,
        * the species of the parents,
        * the health given by `reporduce_helper`,
        * the average of its parents' `growth_rate`, `aging_rate`, `meal_size`, or `success_rate`. 

```python
# Example for Plant reproduction
plant1 = Plant("OakTree1", "Tree", 120, 5)
plant2 = Plant("OakTree2", "Tree", 90, 7)
offspring_plant = plant1.reproduce(plant2, "OakTree3")
print(offspring_plant.get_name())        # Expected: "OakTree3"
print(offspring_plant.get_species())     # Expected: "Tree"
print(offspring_plant.get_health())      # Expected: 35 (120 + 80) / 6
print(offspring_plant.get_growth_rate()) # Expected: 6 (average of 5 and 7)
```

```python
# Example for Herbivore reproduction
herbivore1 = Herbivore("Deer1", "Deer", 90, 2, 10)
herbivore2 = Herbivore("Deer2", "Deer", 60, 3, 15)
offspring_herbivore = herbivore1.reproduce(herbivore2, "Deer3")
print(offspring_herbivore.get_name())        # Expected: "Deer3"
print(offspring_herbivore.get_species())     # Expected: "Deer"
print(offspring_herbivore.get_health())      # Expected: 25 (90 + 60) / 6
print(offspring_herbivore.get_meal_size())   # Expected: 12.5 (average of 10 and 15)
```

```python
# Example for Carnivore Reproduction
carnivore1 = Carnivore("Lion1", "Lion", 100, 2, 0.8)
carnivore2 = Carnivore("Lion2", "Lion", 80, 3, 0.9)
offspring_carnivore = carnivore1.reproduce(carnivore2, "Lion3")
print(offspring_carnivore.get_name())        # Expected: "Lion3"
print(offspring_carnivore.get_species())     # Expected: "Lion"
print(offspring_carnivore.get_health())      # Expected: 30 (100 + 80) / 6
print(offspring_carnivore.get_success_rate())# Expected: 0.85 (average of 0.8 and 0.9)
```

```python
# Example for Omnivore Reproduction
omnivore1 = Omnivore("Bear1", "Bear", 100, 2, 10, 0.75)
omnivore2 = Omnivore("Bear2", "Bear", 80, 3, 15, 0.8)
offspring_omnivore = omnivore1.reproduce(omnivore2, "Bear3")
print(offspring_omnivore.get_name())         # Expected: "Bear3"
print(offspring_omnivore.get_species())      # Expected: "Mammal"
print(offspring_omnivore.get_health())       # Expected: 30 (100 + 80) / 6
print(offspring_omnivore.get_meal_size())    # Expected: 12.5 (average of 10 and 15)
print(offspring_omnivore.get_success_rate()) # Expected: 0.775 (average of 0.75 and 0.8)
```

```python
grader.check("S7")
```

## Section 8: The `Ecosystem` Class
The `Ecosystem` class simulates an environment where plants grow, animals eat or hunt, and organisms reproduce. Organisms are grouped into separate lists based on their type, and actions are taken in a randomized order during each cycle of the simulation.

#### Attributes:
* `plants`: A list that holds all the plants in the ecosystem.
* `herbivores`: A list that holds all the herbivores in the ecosystem.
* `carnivores`: A list that holds all the carnivores in the ecosystem.
* `omnivores`: A list that holds all the omnivores in the ecosystem.

#### Methods:
1. `add_plant(self, plant)`:
    * Adds a plant to the ecosystem by appending it to the plants list.
2. `add_herbivore(self, herbivore)`:
    * Adds a herbivore to the ecosystem by appending it to the herbivores list.
3. `add_carnivore(self, carnivore)`:
    * Adds a carnivore to the ecosystem by appending it to the carnivores list.
4. `add_omnivore(self, omnivore)`:
    * Adds an omnivore to the ecosystem by appending it to the omnivores list.
5. `grow_plants(self)`:
    * Simulates the growth of all plants in the ecosystem by calling the `grow()` method for each plant in the plants list.
6. `remove_dead_organisms(self)`:
    * Removes dead plants and animals from the ecosystem. Any organism that is no longer alive is removed from its respective list.
7. `reproduce_organisms(self, animal_list)` **(already provided for you)**:
    * Neighbors of the same species in the given `animal_list` will reproduce.
    * The new offspring created from reproduction will be added to the respective list (`herbivores`, `carnivores`, or `omnivores`). 
8. `simulate(self)` **(already provided for you)**:
    * Simulates one cycle of the ecosystem following the below steps:
        1. Plant Growth: All plants grow by calling `grow_plants()`.
        2. Animal Actions: All animals (`herbivores`, `carnivores`, and `omnivores`) are shuffled into a random order and take actions:
            * Herbivores: Eat the first available plant in the `plants` list.
            * Carnivores: Hunt the first available animal from the `herbivores` list. If none are available, hunt the first available animal from the `omnivores` list.
            * Omnivores: Eat the first available plant in the `plants` list. If none are available, hunt the first available animal from the `herbivores` list.
        3. Remove Dead Organisms: Dead organisms are removed by calling `remove_dead_organisms()`.
        4. Animal Aging: All animals age.
        5. Reproduction: Organisms reproduce by calling `reproduce_organisms()` given the shuffled animal list created in step 2.
9. `print_organisms(self)`:
    * Lists the current state of the ecosystem by printing all plants, herbivores, carnivores, and omnivores.

```python
class Ecosystem:
    ...
    
    def reproduce_organisms(self, animal_list):
        for i in range(len(animal_list) - 1):
            parent1 = animal_list[i]
            parent2 = animal_list[i + 1]

            if type(parent1) == type(parent2) and parent1.get_species() == parent2.get_species():
                new_name = f"{parent1.get_name()}_{self.name_counter}"
                self.name_counter += 1
                offspring = parent1.reproduce(parent2, new_name)
                if type(offspring) == Plant:
                    self.plants.append(offspring)
                elif type(offspring) == Herbivore:
                    self.herbivores.append(offspring)
                elif type(offspring) == Carnivore:
                    self.carnivores.append(offspring)
                elif type(offspring) == Omnivore:
                    self.omnivores.append(offspring)

    def simulate(self):
        # plants grow
        self.grow_plants()
        
        # animals eat
        all_animals = self.herbivores + self.carnivores + self.omnivores
        random.shuffle(all_animals)

        for animal in all_animals:
            if animal.is_alive(): 
                if type(animal) == Herbivore:
                    for plant in self.plants:
                        if plant.is_alive():
                            animal.eat_plant(plant)
                            break 
    
                elif type(animal) == Carnivore:
                    for prey in self.herbivores + self.omnivores:
                        if prey.is_alive():
                            animal.hunt(prey)
                            break
    
                elif type(animal) == Omnivore:
                    for plant in self.plants:
                        if plant.is_alive():
                            animal.eat_plant(plant)
                            break 
                    else:
                        for prey in self.herbivores:
                            if prey.is_alive():
                                animal.hunt(prey)
                                break  
        # remove dead organisms
        self.remove_dead_organisms()
        
        # animals age
        all_animals = self.herbivores + self.carnivores + self.omnivores
        for animal in all_animals:
            animal.age()
            
        # remove dead organisms
        self.remove_dead_organisms()
        
        # plants reproduce
        all_plants = [x for x in self.plants]
        random.shuffle(all_plants)
        self.reproduce_organisms(all_plants)

        # animals reproduce
        all_animals = self.herbivores + self.carnivores + self.omnivores
        random.shuffle(all_animals)
        self.reproduce_organisms(all_animals)
```

```python
# Example
random.seed(51)
ecosystem = Ecosystem()

# Add plants, herbivores, carnivores, and omnivores to the ecosystem
plant1 = Plant("OakTree", "Tree", 50, 5)
plant2 = Plant("Berries1", "Berry", 30, 3)
plant3 = Plant("Berries2", "Berry", 30, 3)
plant4 = Plant("Berries3", "Berry", 30, 3)
herbivore1 = Herbivore("Deer1", "Deer", 100, 2, 10)
herbivore2 = Herbivore("Deer2", "Deer", 90, 3, 12)
herbivore3 = Herbivore("Deer3", "Deer", 90, 3, 12)
carnivore = Carnivore("Lion", "Lion", 100, 2, 0.75)
omnivore = Omnivore("Bear", "Bear", 90, 2, 10, 0.8)

ecosystem.add_plant(plant1)
ecosystem.add_plant(plant2)
ecosystem.add_plant(plant3)
ecosystem.add_plant(plant4)
ecosystem.add_herbivore(herbivore1)
ecosystem.add_herbivore(herbivore2)
ecosystem.add_herbivore(herbivore3)
ecosystem.add_carnivore(carnivore)
ecosystem.add_omnivore(omnivore)

# Print initial state
print("Initial ecosystem state:")
ecosystem.print_organisms()

# Expected:
# """
# Initial ecosystem state:
# Plants:
# OakTree (species: Tree, health: 50, growth rate: 5)
# Berries1 (species: Berry, health: 30, growth rate: 3)
# Berries2 (species: Berry, health: 30, growth rate: 3)
# Berries3 (species: Berry, health: 30, growth rate: 3)

# Herbivores:
# Deer1 (species: Deer, health: 100, aging rate: 2)
# Deer2 (species: Deer, health: 90, aging rate: 3)
# Deer3 (species: Deer, health: 90, aging rate: 3)

# Carnivores:
# Lion (species: Lion, health: 100, aging rate: 2)

# Omnivores:
# Bear (species: Bear, health: 90, aging rate: 2)
# """


# Simulate one cycle
ecosystem.simulate()

# Print updated state after one cycle
print("\nEcosystem state after one cycle:")
ecosystem.print_organisms()

# Expected:
# """
# Ecosystem state after one cycle:
# Plants:
# OakTree (species: Tree, health: 11, growth rate: 5)
# Berries1 (species: Berry, health: 33, growth rate: 3)
# Berries2 (species: Berry, health: 33, growth rate: 3)
# Berries3 (species: Berry, health: 33, growth rate: 3)

# Herbivores:
# Deer2 (species: Deer, health: 102, aging rate: 3)
# Deer3 (species: Deer, health: 102, aging rate: 3)
# Deer_0 (species: Deer, health: 17.0, aging rate: 2.5)
# Deer_1 (species: Deer, health: 34.0, aging rate: 3.0)

# Carnivores:
# Lion (species: Lion, health: 210, aging rate: 2)

# Omnivores:
# Bear (species: Bear, health: 100, aging rate: 2)
# """
```

```python
# Change print_ecosystem to True to enable printing the ecosystem of the test cases during each cycle. 
print_ecosystem = False
```

```python
grader.check("Q8")
```

In this project, you’ve learned how to simulate an ecosystem using object-oriented programming, inheritance, and randomized interactions. You can now take it further by adding features like dynamic growth and aging rates, flexible meal sizes, and refined hunting success rates to make the simulation more realistic. For a bigger challenge, consider creating a graphical user interface (GUI) to visualize the ecosystem in real-time. These enhancements will not only deepen your understanding but also turn your project into a showcase of your ability to design and implement complex, interactive systems.

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for Q8 (points: 1)
```
>>> cycles = 5
>>> random.seed(368)
>>> ecosystem = Ecosystem()
>>> for i in range(7):
...     ecosystem.add_plant(Plant(f'Berry{i + 1}', 'Berry', random.randint(10, 20), random.randint(2, 5)))
...     ecosystem.add_plant(Plant(f'Tree{i + 1}', 'Tree', random.randint(50, 55), random.randint(3, 5)))
>>> for i in range(4):
...     ecosystem.add_herbivore(Herbivore(f'Deer{i + 1}', 'Deer', random.randint(50, 90), random.randint(10, 15), random.randint(15, 22)))
...     ecosystem.add_herbivore(Herbivore(f'Rabbit{i + 1}', 'Rabbit', random.randint(20, 33), random.randint(5, 10), random.randint(5, 10)))
...     ecosystem.add_herbivore(Herbivore(f'Sheep{i + 1}', 'Sheep', random.randint(50, 70), random.randint(10, 15), random.randint(15, 20)))
>>> for i in range(2):
...     ecosystem.add_carnivore(Carnivore(f'Lion{i + 1}', 'Lion', random.randint(80, 120), random.randint(25, 30), random.uniform(0.6, 0.8)))
>>> for i in range(2):
...     ecosystem.add_omnivore(Omnivore(f'Bear{i + 1}', 'Bear', random.randint(60, 100), random.randint(15, 20), random.randint(12, 15), random.uniform(0.4, 0.6)))
>>> for cycle in range(cycles):
...     if print_ecosystem:
...         print(f'\n--- Cycle {cycle + 1} ---')
...         ecosystem.print_organisms()
...     ecosystem.simulate()
...     for herbivore in ecosystem.herbivores:
...         assert herbivore.get_health() >= 0, f'{herbivore.get_name()} has negative health!'
...     for carnivore in ecosystem.carnivores:
...         assert carnivore.get_health() >= 0, f'{carnivore.get_name()} has negative health!'
...     for omnivore in ecosystem.omnivores:
...         assert omnivore.get_health() >= 0, f'{omnivore.get_name()} has negative health!'
...     for plant in ecosystem.plants:
...         assert plant.get_health() >= 0, f'{plant.get_name()} has negative health!'
...     assert all((plant.is_alive() for plant in ecosystem.plants)), 'There are dead plants remaining in the ecosystem!'
...     assert all((herbivore.is_alive() for herbivore in ecosystem.herbivores)), 'There are dead herbivores remaining in the ecosystem!'
...     assert all((carnivore.is_alive() for carnivore in ecosystem.carnivores)), 'There are dead carnivores remaining in the ecosystem!'
...     assert all((omnivore.is_alive() for omnivore in ecosystem.omnivores)), 'There are dead omnivores remaining in the ecosystem!'
>>> assert len(ecosystem.plants) == 17
>>> assert len(ecosystem.herbivores) == 9
>>> assert len(ecosystem.carnivores) == 5
>>> assert len(ecosystem.omnivores) == 2
>>> if print_ecosystem:
...     print(f'\n--- End ---')
...     ecosystem.print_organisms()
```


Tests for S1 (points: 1)
```
>>> organism_private_test = Organism('Oak1', 'Oak Tree')
>>> assert organism_private_test._Organism__name == 'Oak1'
>>> assert organism_private_test._Organism__species == 'Oak Tree'
>>> assert organism_private_test._health == 100
>>> organism_zero_health = Organism('Rose1', 'Rose', 0)
>>> assert organism_zero_health.get_name() == 'Rose1'
>>> assert organism_zero_health.get_species() == 'Rose'
>>> assert organism_zero_health.get_health() == 0
>>> assert not organism_zero_health.is_alive()
>>> try:
...     organism_zero_health.set_health(50)
... except ValueError as e:
...     assert str(e) == 'Cannot update health for Rose1, organism is dead.'
>>> organism_negative_health = Organism('Tiger1', 'Tiger', -10)
>>> assert organism_negative_health.get_name() == 'Tiger1'
>>> assert organism_negative_health.get_species() == 'Tiger'
>>> assert organism_negative_health.get_health() == -10
>>> assert not organism_negative_health.is_alive()
>>> try:
...     organism_negative_health.set_health(30)
... except ValueError as e:
...     assert str(e) == 'Cannot update health for Tiger1, organism is dead.'
>>> organism_default_health = Organism('Spruce1', 'Spruce Tree')
>>> assert organism_default_health.get_name() == 'Spruce1'
>>> assert organism_default_health.get_species() == 'Spruce Tree'
>>> assert organism_default_health.get_health() == 100
>>> assert organism_default_health.is_alive()
>>> organism_default_health.set_health(75)
>>> assert organism_default_health.get_health() == 75
```


Tests for S2 (points: 1)
```
>>> plant1 = Plant('Oak', 'Tree', 80, 10)
>>> plant2 = Plant('Oak', 'Tree', 80, 10)
>>> plant3 = Plant('Pine', 'Tree', 80, 5)
>>> plant4 = Plant('Oak', 'Tree', 100, 10)
>>> plant5 = Plant('Oak', 'Tree', 100, 5)
>>> plant6 = Plant('Oak', 'Flower', 100, 5)
>>> assert plant1 == plant2
>>> assert plant1 != plant3
>>> assert plant1 != plant4
>>> assert plant4 != plant5
>>> assert plant5 != plant6
>>> initial_health = plant1.get_health()
>>> assert initial_health == 80
>>> plant1.grow()
>>> assert plant1.get_health() == initial_health + 10
>>> dead_plant = Plant('DeadPlant', 'Flower', 0)
>>> try:
...     dead_plant.set_health(20)
... except ValueError as e:
...     assert str(e) == 'Cannot update health for DeadPlant, organism is dead.'
>>> assert str(plant1) == 'Oak (species: Tree, health: 90, growth rate: 10)'
>>> assert repr(plant1) == 'Plant(name=Oak, species=Tree, health=90, growth_rate=10)'
```


Tests for S3 (points: 1)
```
>>> animal1 = Animal('Tiger1', 'Tiger', 90, 3)
>>> initial_health = animal1.get_health()
>>> animal1.age()
>>> assert animal1.get_health() == initial_health - 3
>>> assert str(animal1) == 'Tiger1 (species: Tiger, health: 87, aging rate: 3)'
>>> assert repr(animal1) == 'Animal(name=Tiger1, species=Tiger, health=87, aging_rate=3)'
```


Tests for S4 (points: 1)
```
>>> herbivore1 = Herbivore('Cow1', 'Cow', 90, 2, 10)
>>> plant1 = Plant('Grass1', 'Grass', 25, 5)
>>> initial_plant_health = plant1.get_health()
>>> initial_herbivore_health = herbivore1.get_health()
>>> herbivore1.eat_plant(plant1)
>>> assert herbivore1.get_health() == initial_herbivore_health + 10
>>> assert plant1.get_health() == initial_plant_health - 10
>>> plant1.set_health(7)
>>> initial_herbivore_health = herbivore1.get_health()
>>> herbivore1.eat_plant(plant1)
>>> assert herbivore1.get_health() == initial_herbivore_health + 7
>>> assert plant1.get_health() == 0
>>> initial_herbivore_health = herbivore1.get_health()
>>> herbivore1.eat_plant(plant1)
>>> assert herbivore1.get_health() == initial_herbivore_health
>>> assert plant1.get_health() == 0
```


Tests for S5 (points: 1)
```
>>> random.seed(50)
>>> carnivore1 = Carnivore('Tiger', 'Tiger', 100, 2, 0.8)
>>> prey = Herbivore('Deer', 'Deer', 50, 3)
>>> initial_carnivore_health = carnivore1.get_health()
>>> initial_prey_health = prey.get_health()
>>> carnivore1.hunt(prey)
>>> assert prey.get_health() == 0
>>> assert carnivore1.get_health() == initial_carnivore_health + initial_prey_health
>>> carnivore2 = Carnivore('Lion', 'Lion', 100, 2, 0.2)
>>> prey = Herbivore('Deer', 'Deer', 50, 3)
>>> carnivore2.hunt(prey)
>>> assert prey.get_health() == 50
>>> assert carnivore2.get_health() == 100
>>> dead_prey = Herbivore('Rabbit', 'Rabbit', 0, 3)
>>> carnivore1.hunt(dead_prey)
>>> assert dead_prey.get_health() == 0
>>> assert carnivore2.get_health() == 100
```


Tests for S6 (points: 1)
```
>>> random.seed(40)
>>> omnivore = Omnivore('Bear', 'Bear', 100, 2, 20, 0.8)
>>> plant = Plant('Berries', 'Berries', 25, 5)
>>> prey = Herbivore('Fish', 'Fish', 50, 3)
>>> omnivore.eat_plant(plant)
>>> assert omnivore.get_health() == 120
>>> assert plant.get_health() == 5
>>> omnivore.eat_plant(plant)
>>> assert omnivore.get_health() == 125
>>> assert plant.get_health() == 0
>>> omnivore.hunt(prey)
>>> assert prey.get_health() == 0
>>> assert omnivore.get_health() == 175
>>> omnivore_low_success = Omnivore('Bear', 'Bear', 100, 2, 10, 0.05)
>>> prey = Herbivore('Fish', 'Bear', 50, 3)
>>> omnivore_low_success.hunt(prey)
>>> assert prey.get_health() == 50
>>> assert omnivore_low_success.get_health() == 100
```


Tests for S7 (points: 1)
```
>>> import math
>>> 
>>> def compare_assert(a, b):
...     assert math.isclose(a, b)
>>> plant1 = Plant('OakTree1', 'Tree', 120, 5)
>>> plant2 = Plant('OakTree2', 'Tree', 80, 7)
>>> offspring_plant = plant1.reproduce(plant2, 'OakTree3')
>>> assert offspring_plant.get_name() == 'OakTree3'
>>> assert offspring_plant.get_species() == 'Tree'
>>> compare_assert(offspring_plant.get_health(), (120 + 80) / 6)
>>> compare_assert(offspring_plant.get_growth_rate(), (5 + 7) / 2)
>>> herbivore1 = Herbivore('Deer1', 'Deer', 90, 2, 10)
>>> herbivore2 = Herbivore('Deer2', 'Deer', 80, 3, 15)
>>> offspring_herbivore = herbivore1.reproduce(herbivore2, 'Deer3')
>>> assert offspring_herbivore.get_name() == 'Deer3'
>>> assert offspring_herbivore.get_species() == 'Deer'
>>> compare_assert(offspring_herbivore.get_health(), (90 + 80) / 6)
>>> compare_assert(offspring_herbivore.get_meal_size(), (10 + 15) / 2)
>>> carnivore1 = Carnivore('Lion1', 'Lion', 100, 2, 0.8)
>>> carnivore2 = Carnivore('Lion2', 'Lion', 90, 3, 0.9)
>>> offspring_carnivore = carnivore1.reproduce(carnivore2, 'Lion3')
>>> assert offspring_carnivore.get_name() == 'Lion3'
>>> assert offspring_carnivore.get_species() == 'Lion'
>>> compare_assert(offspring_carnivore.get_health(), (100 + 90) / 6)
>>> compare_assert(offspring_carnivore.get_success_rate(), (0.8 + 0.9) / 2)
>>> omnivore1 = Omnivore('Bear1', 'Bear', 100, 2, 10, 0.75)
>>> omnivore2 = Omnivore('Bear2', 'Bear', 90, 3, 15, 0.8)
>>> offspring_omnivore = omnivore1.reproduce(omnivore2, 'Bear3')
>>> assert offspring_omnivore.get_name() == 'Bear3'
>>> assert offspring_omnivore.get_species() == 'Bear'
>>> compare_assert(offspring_omnivore.get_health(), (100 + 90) / 6)
>>> compare_assert(offspring_omnivore.get_meal_size(), (10 + 15) / 2)
>>> compare_assert(offspring_omnivore.get_success_rate(), (0.75 + 0.8) / 2)
>>> organism1 = Organism('Tree1', 'Tree', 100)
>>> animal1 = Carnivore('Lion', 'Lion', 100, 2, 0.8)
>>> try:
...     organism1.reproduce_helper(animal1)
... except TypeError as e:
...     assert str(e) == 'Can only reproduce with another organism of the same type.'
>>> herbivore1 = Herbivore('Deer', 'Herbivore', 90, 2, 10)
>>> carnivore1 = Carnivore('Lion', 'Lion', 100, 2, 0.8)
>>> try:
...     herbivore1.reproduce(carnivore1, 'DeerLion')
... except TypeError as e:
...     assert str(e) == 'Can only reproduce with another organism of the same type.'
>>> plant1 = Plant('OakTree1', 'Tree', 120, 5)
>>> plant2 = Plant('PineTree', 'Conifer', 100, 6)
>>> try:
...     plant1.reproduce(plant2, 'OakPine')
... except TypeError as e:
...     assert str(e) == 'Can only reproduce with another organism of the same species.'
```
