```python
# Initialize Otter
import otter
grader = otter.Notebook("pa2.ipynb")
```

# Programming Assignment 2 Student Grade Manager

In this assignment, you'll implement a grade manager for this course, CS368: Python for Java Programmers. You’ll work with key data structures such as lists, tuples, sets, and dictionaries to manage student records and calculate grades.

Let’s review the grading components of the course:
* 15 Attendance Questions (5 drops, 1% each, 10% overall)
* 7 Quizzes (1 drop, 5% each, 30% overall)
* 5 Programming Assignments (PAs, no drops, 12% each, 60% overall)
    * You are allowed up to **two** late submissions without penalty, but each late submission must not be more than **5** days overdue.
 
Each student’s record will be stored in a dictionary, with their `student_id` as the key. The value associated with each `student_id` is another dictionary containing the following information. Initially, all records will be set to `None` (Python's version of Java's `null`, indicating that the data is not yet initialized). You will later populate these records with lists representing grades and late days.
```python
student_records = {
    'student_id': {
        'name': 'Student Name',  
        'PA_gradescope_scores': None,  
        'PA_late_days': None,  
        'quiz_grades': None,  
        'attendance_question_grades': None  
    }, 
    'student_id2': {...}
}
```
* `PA_gradescope_scores`: the list will store the auto-graded Gradescope scores for the five programming assignments.
* `PA_late_days`: the list will store how many days late each assignment was submitted (0 if on time).
* `quiz_grades`: the list will store the grades for the seven quizzes.
* `attendance_question_grades`: the list will store the grades for the 15 attendance questions.

## Question 1. Initialize student records
You will be provided with `student_records` where each record entry is initially set to `None`. Your task is to initialize each entry (`PA_gradescope_scores`, `PA_late_days`, `quiz_grades`, `attendance_question_grades`) with a list of zeroes ([0, 0, 0, ...]), where the length of each list corresponds to the number of items for that category:
* 5 for PAs,
* 7 for Quizzes,
* 15 for Attendance Questions.

Make the initialization change directly on the provided `student_records`. There's no return value. 

```python
def initialize_student_records(student_records):
    """
    Initialize the entries in student_records with lists of correct number of zeroes.
    """
    ...
```

```python
# Example input: 
student_records = {
    '12345': {
        'name': 'Alice',
        'PA_gradescope_scores': None,
        'PA_late_days': None,
        'quiz_grades': None,
        'attendance_question_grades': None
    },
    '67890': {
        'name': 'Bob',
        'PA_gradescope_scores': None,
        'PA_late_days': None,
        'quiz_grades': None,
        'attendance_question_grades': None
    }
}
"""
student_records should be changed to
{'12345': {'name': 'Alice',
  'PA_gradescope_scores': [0, 0, 0, 0, 0],
  'PA_late_days': [0, 0, 0, 0, 0],
  'quiz_grades': [0, 0, 0, 0, 0, 0, 0],
  'attendance_question_grades': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]},
 '67890': {'name': 'Bob',
  'PA_gradescope_scores': [0, 0, 0, 0, 0],
  'PA_late_days': [0, 0, 0, 0, 0],
  'quiz_grades': [0, 0, 0, 0, 0, 0, 0],
  'attendance_question_grades': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}}
"""
initialize_student_records(student_records)
student_records
```

```python
grader.check("Q1")
```

## Question 2. Update a single student record entry
Implement a function to update a specific entry in the initialized `student_records`. 

The function will take the following parameters:
* `student_records`: initialized as in Question 1.
* `student_id`: the ID of the student whose record you need to update.
* `assignment_type`: a string specifying the type of assignment to update. The valid values are:
    * 'PA' for Programming Assignment,
    * 'quiz' for Quizzes,
    * 'attendance' for Attendance Questions.
* `assignment_number`: the number of the specific assignment to update (**1-indexed**).
* `new_grade`: the new grade to assign for this specific entry.

Your function should update the corresponding entry in `student_records` based on the `assignment_type` and `assignment_number`.

You should raise a `ValueError` if:
* The `student_id` is not found in student_records. The error message should be: "Student ID <student_id> not found." (This error is already implemented as an example to raising exceptions.)
* The `assignment_type` is invalid. The error message should be: "Invalid assignment type. Valid types are 'PA', 'quiz', or 'attendance'."
* The `assignment_number` is out of bounds for the specified `assignment_type`. 
    * For Programming Assignments, the error message should be: "Assignment number out of bounds for Programming Assignments."
    * For quizzes, the error message should be: "Assignment number out of bounds for Quizzes."
    * For Attendance Questions, the error message should be: "Assignment number out of bounds for Attendance Questions."
  
Hint: documentation on [raising exceptions](https://docs.python.org/3/tutorial/errors.html#raising-exceptions)

```python
def update_record_entry(student_records, student_id, assignment_type, assignment_number, new_grade):
    """
    Update a specific entry in the student's record based on assignment type and assignment number.
    """
    if student_id not in student_records:
        raise ValueError("Student ID not found.")
        
    ...
```

```python
# Example input:
student_records = {
    '12345': {
        'name': 'Alice',
        'PA_gradescope_scores': [85, 90, 78, 88, 92],
        'PA_late_days': [0, 1, 0, 0, 2],
        'PA_grades': [85, 85, 78, 88, 80],
        'quiz_grades': [90, 85, 88, 92, 91, 84, 87],
        'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]
    }
}

update_record_entry(student_records, '12345', 'PA', 3, 95)

# Expected results: 
# student_records['12345']['PA_gradescope_scores'] -> [85, 90, 95, 88, 92]
student_records['12345']['PA_gradescope_scores']
```

```python
grader.check("Q2")
```

## Question 3. Batch update of an assignment
Implement a function that performs batch updates for a specific assignment type and assignment number using a dictionary of updates. 

The function will take the following parameters:
* `student_records`: initialized as in Question 1.
* `assignment_type`: same as Question 2. 
* `assignment_number`: same as Question 2.
* `batch_updates`: a dictionary where each key is a `student_id` and the corresponding value is the new grade to assign for the specified `assignment_type` and `assignment_number`.

You should first raise a ValueError using the same error messages as in Question 2 for:
* An invalid `assignment_type`.
* An out-of-bounds `assignment_number` for the specified `assignment_type`.

Your function should then loop through the `batch_updates` dictionary and update the corresponding grades for each student. 

The function should return any invalid `student_id`s in a set.

```python
def batch_update_grades(student_records, assignment_type, assignment_number, batch_updates):
    """
    Perform batch updates for a specific assignment type and assignment number using a dictionary of updates.
    """
    ...
```

```python
# Example input: 
student_records = {
    '12345': {
        'name': 'Alice',
        'PA_gradescope_scores': [85, 90, 0, 0, 0],
        'PA_late_days': [0, 1, 0, 0, 0],
        'PA_grades': [85, 85, 0, 0, 0],
        'quiz_grades': [90, 85, 0, 0, 0, 0, 0],
        'attendance_question_grades': [1, 1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
    },
    '67890': {
        'name': 'Bob',
        'PA_gradescope_scores': [88, 92, 0, 0, 0],
        'PA_late_days': [0, 0, 0, 0, 0],
        'PA_grades': [88, 92, 0, 0, 0],
        'quiz_grades': [85, 87, 0, 0, 0, 0, 0],
        'attendance_question_grades': [1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
    }
}

batch_updates = {
    '12345': 95,
    '67890': 90,
    '99999': 85  
}

invalid_student_ids = batch_update_grades(student_records, 'PA', 3, batch_updates)

# Expected results:
# student_records['12345']['PA_gradescope_scores'] -> [85, 90, 95, 0, 0]
# student_records['67890']['PA_gradescope_scores'] -> [88, 92, 90, 0, 0]
# invalid_student_ids -> {'99999'}

print(student_records['12345']['PA_gradescope_scores'])
print(student_records['67890']['PA_gradescope_scores'])
print(invalid_student_ids)
```

```python
grader.check("Q3")
```

## Question 4. Compute average with drops
You will implement a function that computes the average grade for quizzes or attendance, considering a certain number of drops. The lowest grades in the list will be dropped.

The function will take the following inputs:

* `grades`: a list of grades (e.g., quiz or attendance question grades).
* `num_drops`: an integer specifying the number of lowest grades to drop before computing the average.

If the number of grades is less than or equal to the number of drops, return 0 as the average.

```python
def compute_average_with_drops(grades, num_drops):
    """
    Compute the average of grades after dropping the lowest num_drops grades.
    """
    ...
```

```python
grades = [90, 85, 88, 92, 91, 84, 87]
num_drops = 1
# Expected output: (90 + 85 + 88 + 92 + 91 + 0 + 87) / 6 = 88.83333333333333
compute_average_with_drops(grades, num_drops)
```

```python
grader.check("Q4")
```

## Question 5. Compute PA average with late submission opportunities
Implement a function to compute the average grade for Programming Assignments (PAs), factoring in late submissions. 

In CS368, students are allowed up to **two** late submissions, with each submission allowed to be up to **5** days late without penalty. 
Any submission that is more than 5 days late is not allowed and will receive a grade of 0. The first two late submissions that is not late for more than 5 days will be allowed, but any late submissions after them will receive a grade of 0.   

In the function that you will implement, the allowed number of late submissions and the max late days are variables given to you.  

The function will take the following parameters:
* `pa_grades`: a list of grades for the Programming Assignments.
* `pa_late_days`: a list of integers representing how many days late each corresponding Programming Assignment was submitted.
* `allowed_late_submissions`: an integer specifying how many late submissions are allowed without penalty.
* `max_late_days`: an integer specifying the maximum number of late days allowed without penalty for each assignment.

```python
def compute_pa_average(pa_grades, pa_late_days, allowed_late_submissions, max_late_days):
    """
    Compute the average PA score, factoring in late submission penalties.
    """
    ...
```

```python
pa_grades = [85, 90, 78, 88, 92]
pa_late_days = [0, 6, 2, 3, 3]
allowed_late_submissions = 2
max_late_days = 5
# Expected output: (85 + 0 + 78 + 88 + 0) / 5 = 50.2
compute_pa_average(pa_grades, pa_late_days, allowed_late_submissions, max_late_days)
```

```python
grader.check("Q5")
```

## Question 6. Compute final PA, Quiz, and Attendance grades from complete records
Implement a function that processes each student's complete records and adds their final grades for Programming Assignments, Quizzes, and Attendance to their record.

The function will take the following parameters:

* `student_records`: all `PA_gradescope_scores`, `PA_late_days`, `quiz_grades`, and `attendance_question_grades` records in `student_records` are filled.
* `allowed_late_submissions`: the number of late submissions allowed for PAs without penalty.
* `max_late_days`: the maximum number of late days allowed without penalty per PA.
* `quiz_drops`: the number of lowest quiz grades to drop.
* `attendance_drops`: the number of lowest attendance grades to drop.

Use `compute_average_with_drops` and `compute_pa_average` functions you implemented in Question 4 and 5 to compute the final PA, quiz, and attendance grades and add them back to each student’s record under the keys:
* `final_PA_grade`
* `final_quiz_grade`
* `final_attendance_grade`

```python
def update_final_assignment_grades(student_records, allowed_late_submissions, max_late_days, quiz_drops, attendance_drops):
    """
    Compute and add final PA, quiz, and attendance grades to each student's record.
    """
    ...
```

```python
# Example input: 
student_records = {
    '12345': {
        'PA_gradescope_scores': [85, 90, 78, 88, 92],
        'PA_late_days': [0, 6, 0, 0, 3],
        'quiz_grades': [90, 85, 88, 92, 91, 84, 87],
        'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]
    }
}
"""
Expected result of student_records: 
{'12345': {'PA_gradescope_scores': [85, 90, 78, 88, 92],
  'PA_late_days': [0, 6, 0, 0, 3],
  'quiz_grades': [90, 85, 88, 92, 91, 84, 87],
  'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
  'final_PA_grade': 68.6,
  'final_quiz_grade': 88.83333333333333,
  'final_attendance_grade': 1.0}}
"""
update_final_assignment_grades(student_records, allowed_late_submissions=2, max_late_days=5, quiz_drops=1, attendance_drops=5)
student_records
```

```python
grader.check("Q6")
```

## Question 7. Compute final grade and letter grade
Implement a function to compute the final grade and letter grade for each student based on the weighted grades for Programming Assignments (PAs), quizzes, and attendance.

The function will take the following inputs:

* `student_records`: all `final_PA_grade`, `final_quiz_grade`, and `final_attendance_grade` are filled in `student_records` as in Question 6. 
* `pa_weight`: the weight of PAs in the final grade.
* `quiz_weight`: the weight of quizzes in the final grade.
* `attendance_weight`: the weight of attendance in the final grade.

Your function should:
* Compute the final grade for each student.
* Assign a letter grade of 'P' if the final grade is 70 or higher, otherwise 'NP'.
* Add `final_grade` and `letter_grade` to each student’s record.
* Return a list of tuples containing the `student_id` and `letter_grade`.

```python
def compute_final_grades(student_records, pa_weight, quiz_weight, attendance_weight):
    """
    Compute the final grade and letter grade for each student, and return a list of tuples with student_id and letter_grade.
    """
    ...
```

```python
# Example input:
student_records = {
    '12345': {
        'final_PA_grade': 86,
        'final_quiz_grade': 90,
        'final_attendance_grade': 95
    },
    '67890': {
        'final_PA_grade': 60,
        'final_quiz_grade': 65,
        'final_attendance_grade': 70
    }
}

# Expected output: [('12345', 'P'), ('67890', 'NP')]
compute_final_grades(student_records, pa_weight=0.6, quiz_weight=0.3, attendance_weight=0.1)
```

```python
student_records
```

```python
grader.check("Q7")
```

## Question 8. Putting it all together
You will implement a function that, given `student_records` with all PA, quiz, and attendance data filled, and a dictionary of configurations, computes both the final grades for each assignment type (PA, quiz, attendance) and the overall final grade, as well as assigns a letter grade using functions in Question 6 and 7. 

The function will take the following parameters:

* `student_records`: all `PA_gradescope_scores`, `PA_late_days`, `quiz_grades`, and `attendance_question_grades` records in `student_records` are filled. 
* `configuration`: A dictionary specifying:
    * `allowed_late_submissions`
    * `max_late_days`
    * `quiz_drops`
    * `attendance_drops`
    * `pa_weight`
    * `quiz_weight`
    * `attendance_weight`

Your function should achieve what have been done in question 6 and 7. 

```python
def put_it_together(student_records, configuration):
    """
    Compute the final assignment grades (PA, quiz, attendance), the final overall grade, 
    and return the letter grade for each student.
    """

    ...
```

```python
# Example input: 
student_records = {
    '12345': {
        'PA_gradescope_scores': [85, 90, 78, 88, 92],
        'PA_late_days': [0, 5, 0, 0, 3],
        'quiz_grades': [90, 85, 88, 92, 91, 84, 87],
        'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]
    }
}

configuration = {
    'allowed_late_submissions': 2,
    'max_late_days': 5,
    'quiz_drops': 1,
    'attendance_drops': 5,
    'pa_weight': 0.6,
    'quiz_weight': 0.3,
    'attendance_weight': 0.1
}

put_it_together(student_records, configuration)
```

```python
student_records
```

```python
grader.check("Q8")
```

Now you know how to build a grade manager using a variety of data structures! You’ve worked with dictionaries to store student records, lists for grades and attendance, sets for handling unique values, and tuples for efficiently managing immutable data like final grade summaries. By combining these structures, you've streamlined the process of updating records, computing grades, and handling late submissions. With these skills, your grade manager is ready to handle complex student data and grading tasks with ease!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for Q1 (points: 1)
```
>>> test_records = {'11111': {'name': 'Dana', 'PA_gradescope_scores': None, 'PA_late_days': None, 'quiz_grades': None, 'attendance_question_grades': None}, '22222': {'name': 'Eli', 'PA_gradescope_scores': None, 'PA_late_days': None, 'quiz_grades': None, 'attendance_question_grades': None}}
>>> initialize_student_records(test_records)
>>> assert test_records['11111']['PA_gradescope_scores'] == [0, 0, 0, 0, 0]
>>> assert test_records['11111']['PA_late_days'] == [0, 0, 0, 0, 0]
>>> assert test_records['11111']['quiz_grades'] == [0, 0, 0, 0, 0, 0, 0]
>>> assert test_records['11111']['attendance_question_grades'] == [0] * 15
>>> assert test_records['22222']['PA_gradescope_scores'] == [0, 0, 0, 0, 0]
>>> assert test_records['22222']['PA_late_days'] == [0, 0, 0, 0, 0]
>>> assert test_records['22222']['quiz_grades'] == [0, 0, 0, 0, 0, 0, 0]
>>> assert test_records['22222']['attendance_question_grades'] == [0] * 15
```


Tests for Q2 (points: 1)
```
>>> test_records = {'12345': {'name': 'Alice', 'PA_gradescope_scores': [85, 90, 78, 88, 92], 'PA_late_days': [0, 1, 0, 0, 2], 'PA_grades': [85, 85, 78, 88, 80], 'quiz_grades': [90, 85, 88, 92, 91, 84, 87], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> update_record_entry(test_records, '12345', 'PA', 3, 95)
>>> assert test_records['12345']['PA_gradescope_scores'] == [85, 90, 95, 88, 92]
```
```
>>> test_records = {'12345': {'name': 'Alice', 'PA_gradescope_scores': [85, 90, 78, 88, 92], 'PA_late_days': [0, 1, 0, 0, 2], 'PA_grades': [85, 85, 78, 88, 80], 'quiz_grades': [90, 85, 88, 92, 91, 84, 87], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> update_record_entry(test_records, '12345', 'quiz', 5, 94)
>>> assert test_records['12345']['quiz_grades'][4] == 94
```
```
>>> test_records = {'12345': {'name': 'Alice', 'PA_gradescope_scores': [85, 90, 78, 88, 92], 'PA_late_days': [0, 1, 0, 0, 2], 'PA_grades': [85, 85, 78, 88, 80], 'quiz_grades': [90, 85, 88, 92, 91, 84, 87], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> update_record_entry(test_records, '12345', 'attendance', 5, 1)
>>> assert test_records['12345']['attendance_question_grades'][4] == 1
```
```
>>> test_records = {'12345': {'name': 'Alice', 'PA_gradescope_scores': [85, 90, 78, 88, 92], 'PA_late_days': [0, 1, 0, 0, 2], 'PA_grades': [85, 85, 78, 88, 80], 'quiz_grades': [90, 85, 88, 92, 91, 84, 87], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> try:
...     update_record_entry(test_records, '99999', 'PA', 3, 95)
... except ValueError as e:
...     assert str(e) == 'Student ID not found.'
```
```
>>> test_records = {'12345': {'name': 'Alice', 'PA_gradescope_scores': [85, 90, 78, 88, 92], 'PA_late_days': [0, 1, 0, 0, 2], 'PA_grades': [85, 85, 78, 88, 80], 'quiz_grades': [90, 85, 88, 92, 91, 84, 87], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> try:
...     update_record_entry(test_records, '12345', 'project', 1, 95)
... except ValueError as e:
...     assert str(e) == "Invalid assignment type. Valid types are 'PA', 'quiz', or 'attendance'."
```
```
>>> test_records = {'12345': {'name': 'Alice', 'PA_gradescope_scores': [85, 90, 78, 88, 92], 'PA_late_days': [0, 1, 0, 0, 2], 'PA_grades': [85, 85, 78, 88, 80], 'quiz_grades': [90, 85, 88, 92, 91, 84, 87], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> try:
...     update_record_entry(test_records, '12345', 'quiz', 8, 95)
... except ValueError as e:
...     assert str(e) == 'Assignment number out of bounds for Quizzes.'
```
```
>>> test_records = {'12345': {'name': 'Alice', 'PA_gradescope_scores': [85, 90, 78, 88, 92], 'PA_late_days': [0, 1, 0, 0, 2], 'PA_grades': [85, 85, 78, 88, 80], 'quiz_grades': [90, 85, 88, 92, 91, 84, 87], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> try:
...     update_record_entry(test_records, '12345', 'PA', -1, 95)
... except ValueError as e:
...     assert str(e) == 'Assignment number out of bounds for Programming Assignments.'
```


Tests for Q3 (points: 1)
```
>>> student_records = {'12345': {'name': 'Alice', 'PA_gradescope_scores': [85, 90, 0, 0, 0], 'PA_late_days': [0, 1, 0, 0, 0], 'PA_grades': [85, 85, 0, 0, 0], 'quiz_grades': [90, 85, 0, 0, 0, 0, 0], 'attendance_question_grades': [1, 1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, '67890': {'name': 'Bob', 'PA_gradescope_scores': [88, 92, 0, 0, 0], 'PA_late_days': [0, 0, 0, 0, 0], 'PA_grades': [88, 92, 0, 0, 0], 'quiz_grades': [85, 87, 0, 0, 0, 0, 0], 'attendance_question_grades': [1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}}
>>> batch_updates = {'12345': 95, '67890': 90}
>>> invalid_student_ids = batch_update_grades(student_records, 'PA', 3, batch_updates)
>>> assert student_records['12345']['PA_gradescope_scores'][2] == 95
>>> assert student_records['67890']['PA_gradescope_scores'][2] == 90
>>> assert invalid_student_ids == set()
```
```
>>> student_records = {'12345': {'name': 'Alice', 'PA_gradescope_scores': [85, 90, 0, 0, 0], 'PA_late_days': [0, 1, 0, 0, 0], 'PA_grades': [85, 85, 0, 0, 0], 'quiz_grades': [90, 85, 0, 0, 0, 0, 0], 'attendance_question_grades': [1, 1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, '67890': {'name': 'Bob', 'PA_gradescope_scores': [88, 92, 0, 0, 0], 'PA_late_days': [0, 0, 0, 0, 0], 'PA_grades': [88, 92, 0, 0, 0], 'quiz_grades': [85, 87, 0, 0, 0, 0, 0], 'attendance_question_grades': [1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}}
>>> batch_updates = {'12345': 70, '67890': 80, '99999': 85, 'abcd': 77}
>>> invalid_student_ids = batch_update_grades(student_records, 'PA', 3, batch_updates)
>>> assert student_records['12345']['PA_gradescope_scores'][2] == 70
>>> assert student_records['67890']['PA_gradescope_scores'][2] == 80
>>> assert invalid_student_ids == {'99999', 'abcd'}
```
```
>>> try:
...     batch_update_grades(student_records, 'project', 3, batch_updates)
... except ValueError as e:
...     assert str(e) == "Invalid assignment type. Valid types are 'PA', 'quiz', or 'attendance'."
```
```
>>> try:
...     batch_update_grades(student_records, 'PA', 6, batch_updates)
... except ValueError as e:
...     assert str(e) == 'Assignment number out of bounds for Programming Assignments.'
```


Tests for Q4 (points: 1)
```
>>> grades = [90, 85, 88, 92, 91, 84, 87]
>>> num_drops = 1
>>> average = compute_average_with_drops(grades, num_drops)
>>> assert average == (90 + 85 + 88 + 92 + 91 + 87) / 6
```
```
>>> grades = [1, 1, 1, 1, 0, 1, 1, 0, 1, 0, 1, 1, 1, 1, 1]
>>> num_drops = 5
>>> average = compute_average_with_drops(grades, num_drops)
>>> assert average == (1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1) / 10
```
```
>>> grades = [50, 60]
>>> num_drops = 1
>>> average = compute_average_with_drops(grades, num_drops)
>>> assert average == 60
```
```
>>> grades = [70]
>>> num_drops = 1
>>> average = compute_average_with_drops(grades, num_drops)
>>> assert average == 0
```
```
>>> grades = [100, 100, 100, 100]
>>> num_drops = 10
>>> average = compute_average_with_drops(grades, num_drops)
>>> assert average == 0
```
```
>>> grades = [90, 85, 88, 92, 91, 84, 87]
>>> num_drops = 0
>>> average = compute_average_with_drops(grades, num_drops)
>>> assert average == sum(grades) / len(grades)
```


Tests for Q5 (points: 1)
```
>>> pa_grades = [85, 90, 78, 88, 92]
>>> pa_late_days = [0, 4, 0, 0, 3]
>>> allowed_late_submissions = 2
>>> max_late_days = 5
>>> average = compute_pa_average(pa_grades, pa_late_days, allowed_late_submissions, max_late_days)
>>> assert average == (85 + 90 + 78 + 88 + 92) / 5
```
```
>>> pa_grades = [85, 90, 78, 88, 92]
>>> pa_late_days = [0, 6, 2, 3, 3]
>>> allowed_late_submissions = 2
>>> max_late_days = 5
>>> average = compute_pa_average(pa_grades, pa_late_days, allowed_late_submissions, max_late_days)
>>> assert average == (85 + 0 + 78 + 88 + 0) / 5
```
```
>>> pa_grades = [85, 90, 78, 88, 92]
>>> pa_late_days = [0, 6, 0, 0, 3]
>>> allowed_late_submissions = 2
>>> max_late_days = 5
>>> average = compute_pa_average(pa_grades, pa_late_days, allowed_late_submissions, max_late_days)
>>> assert average == (85 + 0 + 78 + 88 + 92) / 5
```
```
>>> pa_grades = [85, 90, 78, 88, 92]
>>> pa_late_days = [0, 3, 0, 2, 5]
>>> allowed_late_submissions = 2
>>> max_late_days = 5
>>> average = compute_pa_average(pa_grades, pa_late_days, allowed_late_submissions, max_late_days)
>>> assert average == (85 + 90 + 78 + 88 + 0) / 5
```
```
>>> pa_grades = [85, 90, 78, 88, 92]
>>> pa_late_days = [6, 7, 3, 4, 8]
>>> allowed_late_submissions = 2
>>> max_late_days = 5
>>> average = compute_pa_average(pa_grades, pa_late_days, allowed_late_submissions, max_late_days)
>>> assert average == (78 + 88) / 5
```


Tests for Q6 (points: 1)
```
>>> student_records = {'12345': {'PA_gradescope_scores': [85, 90, 78, 88, 92], 'PA_late_days': [0, 6, 0, 0, 3], 'quiz_grades': [90, 85, 88, 92, 91, 84, 87], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> update_final_assignment_grades(student_records, allowed_late_submissions=2, max_late_days=5, quiz_drops=1, attendance_drops=5)
>>> assert student_records['12345']['final_PA_grade'] == (85 + 0 + 78 + 88 + 92) / 5
>>> assert student_records['12345']['final_quiz_grade'] == (90 + 85 + 88 + 92 + 91 + 87) / 6
>>> assert student_records['12345']['final_attendance_grade'] == (1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1) / 10
```
```
>>> student_records = {'67890': {'PA_gradescope_scores': [95, 85, 80, 90, 100], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [90, 85, 88, 92, 87, 85, 89], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> update_final_assignment_grades(student_records, allowed_late_submissions=2, max_late_days=5, quiz_drops=0, attendance_drops=0)
>>> assert student_records['67890']['final_PA_grade'] == (95 + 85 + 80 + 90 + 100) / 5
>>> assert student_records['67890']['final_quiz_grade'] == (90 + 85 + 88 + 92 + 87 + 85 + 89) / 7
>>> assert student_records['67890']['final_attendance_grade'] == 1.0
```
```
>>> student_records = {'24680': {'PA_gradescope_scores': [70, 75, 80, 85, 90], 'PA_late_days': [6, 6, 7, 8, 9], 'quiz_grades': [85, 80, 78, 90, 95, 88, 91], 'attendance_question_grades': [1, 1, 0, 1, 1, 0, 1, 1, 0, 0, 1, 0, 1, 1, 0]}}
>>> update_final_assignment_grades(student_records, allowed_late_submissions=2, max_late_days=5, quiz_drops=1, attendance_drops=5)
>>> assert student_records['24680']['final_PA_grade'] == 0
>>> assert student_records['24680']['final_quiz_grade'] == (85 + 80 + 0 + 90 + 95 + 88 + 91) / 6
>>> assert student_records['24680']['final_attendance_grade'] == (1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 0) / 10
```
```
>>> student_records = {'54321': {'PA_gradescope_scores': [95, 85, 80, 90, 100], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [90, 0, 0, 0, 0, 0, 0], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}}
>>> update_final_assignment_grades(student_records, allowed_late_submissions=2, max_late_days=5, quiz_drops=2, attendance_drops=5)
>>> assert student_records['54321']['final_PA_grade'] == (95 + 85 + 80 + 90 + 100) / 5
>>> assert student_records['54321']['final_quiz_grade'] == 90 / 5
>>> assert student_records['54321']['final_attendance_grade'] == (1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1) / 10
```
```
>>> student_records = {'12345': {'PA_gradescope_scores': [85, 90, 78, 88, 92], 'PA_late_days': [0, 6, 0, 0, 3], 'quiz_grades': [90, 85, 88, 92, 91, 84, 87], 'attendance_question_grades': [1, 0, 1, 1, 0, 1, 0, 1, 1, 0, 1, 0, 1, 1, 1]}, '67890': {'PA_gradescope_scores': [95, 85, 80, 90, 100], 'PA_late_days': [0, 4, 0, 6, 4], 'quiz_grades': [90, 0, 88, 92, 87, 0, 89], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1]}, '24680': {'PA_gradescope_scores': [70, 75, 80, 85, 90], 'PA_late_days': [0, 0, 3, 4, 4], 'quiz_grades': [85, 80, 78, 90, 95, 88, 91], 'attendance_question_grades': [0, 1, 0, 1, 1, 0, 1, 1, 0, 1, 0, 1, 0, 0, 0]}}
>>> allowed_late_submissions = 2
>>> max_late_days = 5
>>> quiz_drops = 1
>>> attendance_drops = 5
>>> update_final_assignment_grades(student_records, allowed_late_submissions, max_late_days, quiz_drops, attendance_drops)
>>> assert student_records['12345']['final_PA_grade'] == (85 + 0 + 78 + 88 + 92) / 5
>>> assert student_records['12345']['final_quiz_grade'] == (90 + 85 + 88 + 92 + 91 + 0 + 87) / 6
>>> assert student_records['12345']['final_attendance_grade'] == (1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1) / 10
>>> assert student_records['67890']['final_PA_grade'] == (95 + 85 + 80 + 0 + 100) / 5
>>> assert student_records['67890']['final_quiz_grade'] == (90 + 0 + 88 + 92 + 87 + 0 + 89) / 6
>>> assert student_records['67890']['final_attendance_grade'] == (1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1 + 1) / 10
>>> assert student_records['24680']['final_PA_grade'] == (70 + 75 + 80 + 85 + 0) / 5
>>> assert student_records['24680']['final_quiz_grade'] == (85 + 80 + 0 + 90 + 95 + 88 + 91) / 6
>>> assert student_records['24680']['final_attendance_grade'] == (1 + 1 + 1 + 1 + 1 + 1 + 1 + 0 + 0 + 0) / 10
```


Tests for Q7 (points: 1)
```
>>> student_records = {'12345': {'final_PA_grade': 86, 'final_quiz_grade': 90, 'final_attendance_grade': 95}, '67890': {'final_PA_grade': 60, 'final_quiz_grade': 65, 'final_attendance_grade': 70}, '11111': {'final_PA_grade': 70, 'final_quiz_grade': 70, 'final_attendance_grade': 70}}
>>> result = compute_final_grades(student_records, pa_weight=0.6, quiz_weight=0.3, attendance_weight=0.1)
>>> assert student_records['12345']['final_grade'] == 86 * 0.6 + 90 * 0.3 + 95 * 0.1
>>> assert student_records['12345']['letter_grade'] == 'P'
>>> assert student_records['67890']['final_grade'] == 60 * 0.6 + 65 * 0.3 + 70 * 0.1
>>> assert student_records['67890']['letter_grade'] == 'NP'
>>> assert student_records['11111']['final_grade'] == 70
>>> assert student_records['11111']['letter_grade'] == 'P'
>>> assert result == [('12345', 'P'), ('67890', 'NP'), ('11111', 'P')]
```


Tests for Q8 (points: 1)
```
>>> import random
>>> 
>>> def generate_random_grades(num_grades, min_grade=70, max_grade=100):
...     return [random.randint(min_grade, max_grade) for _ in range(num_grades)]
>>> 
>>> def generate_random_late_days(num_days, max_late_days=10, prob_on_time=0.8, prob_small_late=0.15):
...     late_days = []
...     for _ in range(num_days):
...         if random.random() < prob_on_time:
...             late_days.append(0)
...         elif random.random() < prob_small_late:
...             late_days.append(random.randint(1, 5))
...         else:
...             late_days.append(random.randint(6, max_late_days))
...     return late_days
>>> 
>>> def generate_random_attendance(num_questions, prob_missed=0.2):
...     attendance = []
...     for _ in range(num_questions):
...         if random.random() < prob_missed:
...             attendance.append(0)
...         else:
...             attendance.append(1)
...     return attendance
>>> 
>>> def create_student_record(num_pa=5, num_quizzes=7, num_attendance=15):
...     return {'PA_gradescope_scores': generate_random_grades(num_pa), 'PA_late_days': generate_random_late_days(num_pa), 'quiz_grades': generate_random_grades(num_quizzes), 'attendance_question_grades': generate_random_attendance(num_attendance)}
>>> 
>>> def create_student_records(num_students):
...     student_records = {}
...     for i in range(1, num_students + 1):
...         student_id = str(10000 + i)
...         student_records[student_id] = create_student_record()
...     return student_records
```
```
>>> test_records = {'10001': {'PA_gradescope_scores': [85, 99, 97, 96, 81], 'PA_late_days': [7, 0, 0, 0, 0], 'quiz_grades': [89, 83, 98, 73, 89, 86, 74], 'attendance_question_grades': [1, 0, 1, 1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1]}, '10002': {'PA_gradescope_scores': [74, 97, 78, 78, 70], 'PA_late_days': [8, 10, 1, 0, 0], 'quiz_grades': [89, 90, 70, 71, 71, 92, 70], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 0, 1]}, '10003': {'PA_gradescope_scores': [100, 87, 85, 75, 71], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [77, 88, 87, 90, 94, 93, 79], 'attendance_question_grades': [1, 1, 1, 0, 1, 1, 1, 0, 0, 0, 1, 0, 1, 0, 0]}, '10004': {'PA_gradescope_scores': [95, 73, 80, 98, 80], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [84, 71, 98, 84, 90, 87, 79], 'attendance_question_grades': [1, 0, 1, 1, 0, 1, 1, 1, 1, 1, 1, 0, 1, 0, 1]}, '10005': {'PA_gradescope_scores': [77, 70, 99, 81, 92], 'PA_late_days': [0, 10, 7, 0, 0], 'quiz_grades': [90, 83, 91, 72, 86, 98, 87], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1]}, '10006': {'PA_gradescope_scores': [100, 86, 94, 77, 93], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [73, 90, 100, 89, 72, 97, 81], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 1, 1, 1, 1]}, '10007': {'PA_gradescope_scores': [100, 84, 94, 93, 85], 'PA_late_days': [0, 0, 0, 0, 3], 'quiz_grades': [79, 72, 100, 92, 88, 71, 80], 'attendance_question_grades': [0, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}, '10008': {'PA_gradescope_scores': [99, 92, 71, 79, 86], 'PA_late_days': [0, 0, 10, 0, 0], 'quiz_grades': [96, 96, 98, 74, 70, 72, 79], 'attendance_question_grades': [0, 1, 0, 1, 0, 1, 1, 0, 1, 1, 0, 1, 0, 1, 1]}, '10009': {'PA_gradescope_scores': [80, 82, 82, 70, 78], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [76, 89, 75, 100, 80, 78, 98], 'attendance_question_grades': [1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 0, 1]}, '10010': {'PA_gradescope_scores': [90, 88, 76, 86, 85], 'PA_late_days': [2, 0, 0, 0, 0], 'quiz_grades': [93, 70, 84, 89, 89, 96, 93], 'attendance_question_grades': [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 1, 1, 1, 0, 1]}, '10011': {'PA_gradescope_scores': [80, 82, 88, 80, 75], 'PA_late_days': [0, 6, 0, 0, 0], 'quiz_grades': [90, 89, 99, 80, 71, 87, 92], 'attendance_question_grades': [0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0]}, '10012': {'PA_gradescope_scores': [80, 82, 88, 74, 93], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [80, 99, 72, 70, 86, 75, 73], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 0, 1, 1, 0, 0, 1, 1, 1, 1]}, '10013': {'PA_gradescope_scores': [74, 97, 86, 90, 97], 'PA_late_days': [0, 9, 0, 0, 7], 'quiz_grades': [78, 96, 80, 77, 74, 99, 76], 'attendance_question_grades': [1, 1, 0, 1, 1, 1, 1, 1, 1, 0, 1, 1, 0, 1, 1]}, '10014': {'PA_gradescope_scores': [92, 76, 91, 100, 79], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [75, 70, 78, 89, 76, 93, 83], 'attendance_question_grades': [0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 1, 0]}, '10015': {'PA_gradescope_scores': [87, 100, 99, 96, 78], 'PA_late_days': [0, 7, 0, 0, 0], 'quiz_grades': [83, 73, 74, 83, 97, 98, 99], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}, '10016': {'PA_gradescope_scores': [93, 82, 84, 73, 73], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [92, 81, 97, 97, 90, 98, 75], 'attendance_question_grades': [1, 1, 1, 0, 1, 0, 1, 1, 1, 1, 0, 1, 1, 1, 1]}, '10017': {'PA_gradescope_scores': [89, 99, 82, 74, 82], 'PA_late_days': [0, 6, 9, 6, 0], 'quiz_grades': [84, 76, 71, 77, 83, 70, 100], 'attendance_question_grades': [1, 0, 1, 0, 0, 1, 1, 1, 1, 1, 0, 1, 1, 0, 1]}, '10018': {'PA_gradescope_scores': [71, 94, 93, 80, 86], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [97, 80, 94, 95, 78, 73, 95], 'attendance_question_grades': [1, 1, 1, 0, 1, 0, 1, 1, 1, 0, 1, 1, 1, 1, 0]}, '10019': {'PA_gradescope_scores': [78, 77, 98, 88, 84], 'PA_late_days': [7, 0, 0, 0, 0], 'quiz_grades': [80, 74, 80, 90, 74, 70, 83], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 0, 1, 1, 1, 0, 0, 0, 0, 1]}, '10020': {'PA_gradescope_scores': [83, 87, 76, 84, 71], 'PA_late_days': [0, 0, 10, 0, 0], 'quiz_grades': [78, 71, 81, 95, 91, 99, 81], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1]}}
>>> assert put_it_together(test_records, configuration) == [('10001', 'P'), ('10002', 'NP'), ('10003', 'P'), ('10004', 'P'), ('10005', 'NP'), ('10006', 'P'), ('10007', 'P'), ('10008', 'NP'), ('10009', 'P'), ('10010', 'P'), ('10011', 'NP'), ('10012', 'P'), ('10013', 'NP'), ('10014', 'P'), ('10015', 'NP'), ('10016', 'P'), ('10017', 'NP'), ('10018', 'P'), ('10019', 'NP'), ('10020', 'NP')]
>>> assert test_records == {'10001': {'PA_gradescope_scores': [85, 99, 97, 96, 81], 'PA_late_days': [7, 0, 0, 0, 0], 'quiz_grades': [89, 83, 98, 73, 89, 86, 74], 'attendance_question_grades': [1, 0, 1, 1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1], 'final_PA_grade': 74.6, 'final_quiz_grade': 86.5, 'final_attendance_grade': 1.0, 'final_grade': 70.80999999999999, 'letter_grade': 'P'}, '10002': {'PA_gradescope_scores': [74, 97, 78, 78, 70], 'PA_late_days': [8, 10, 1, 0, 0], 'quiz_grades': [89, 90, 70, 71, 71, 92, 70], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 0, 1], 'final_PA_grade': 45.2, 'final_quiz_grade': 80.5, 'final_attendance_grade': 1.0, 'final_grade': 51.37, 'letter_grade': 'NP'}, '10003': {'PA_gradescope_scores': [100, 87, 85, 75, 71], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [77, 88, 87, 90, 94, 93, 79], 'attendance_question_grades': [1, 1, 1, 0, 1, 1, 1, 0, 0, 0, 1, 0, 1, 0, 0], 'final_PA_grade': 83.6, 'final_quiz_grade': 88.5, 'final_attendance_grade': 0.8, 'final_grade': 76.78999999999999, 'letter_grade': 'P'}, '10004': {'PA_gradescope_scores': [95, 73, 80, 98, 80], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [84, 71, 98, 84, 90, 87, 79], 'attendance_question_grades': [1, 0, 1, 1, 0, 1, 1, 1, 1, 1, 1, 0, 1, 0, 1], 'final_PA_grade': 85.2, 'final_quiz_grade': 87.0, 'final_attendance_grade': 1.0, 'final_grade': 77.32, 'letter_grade': 'P'}, '10005': {'PA_gradescope_scores': [77, 70, 99, 81, 92], 'PA_late_days': [0, 10, 7, 0, 0], 'quiz_grades': [90, 83, 91, 72, 86, 98, 87], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1], 'final_PA_grade': 50.0, 'final_quiz_grade': 89.16666666666667, 'final_attendance_grade': 1.0, 'final_grade': 56.85, 'letter_grade': 'NP'}, '10006': {'PA_gradescope_scores': [100, 86, 94, 77, 93], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [73, 90, 100, 89, 72, 97, 81], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 1, 1, 1, 1], 'final_PA_grade': 90.0, 'final_quiz_grade': 88.33333333333333, 'final_attendance_grade': 1.0, 'final_grade': 80.6, 'letter_grade': 'P'}, '10007': {'PA_gradescope_scores': [100, 84, 94, 93, 85], 'PA_late_days': [0, 0, 0, 0, 3], 'quiz_grades': [79, 72, 100, 92, 88, 71, 80], 'attendance_question_grades': [0, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1], 'final_PA_grade': 91.2, 'final_quiz_grade': 85.16666666666667, 'final_attendance_grade': 1.0, 'final_grade': 80.36999999999999, 'letter_grade': 'P'}, '10008': {'PA_gradescope_scores': [99, 92, 71, 79, 86], 'PA_late_days': [0, 0, 10, 0, 0], 'quiz_grades': [96, 96, 98, 74, 70, 72, 79], 'attendance_question_grades': [0, 1, 0, 1, 0, 1, 1, 0, 1, 1, 0, 1, 0, 1, 1], 'final_PA_grade': 71.2, 'final_quiz_grade': 85.83333333333333, 'final_attendance_grade': 0.9, 'final_grade': 68.56, 'letter_grade': 'NP'}, '10009': {'PA_gradescope_scores': [80, 82, 82, 70, 78], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [76, 89, 75, 100, 80, 78, 98], 'attendance_question_grades': [1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 0, 1], 'final_PA_grade': 78.4, 'final_quiz_grade': 86.83333333333333, 'final_attendance_grade': 1.0, 'final_grade': 73.19, 'letter_grade': 'P'}, '10010': {'PA_gradescope_scores': [90, 88, 76, 86, 85], 'PA_late_days': [2, 0, 0, 0, 0], 'quiz_grades': [93, 70, 84, 89, 89, 96, 93], 'attendance_question_grades': [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 1, 1, 1, 0, 1], 'final_PA_grade': 85.0, 'final_quiz_grade': 90.66666666666667, 'final_attendance_grade': 0.9, 'final_grade': 78.29, 'letter_grade': 'P'}, '10011': {'PA_gradescope_scores': [80, 82, 88, 80, 75], 'PA_late_days': [0, 6, 0, 0, 0], 'quiz_grades': [90, 89, 99, 80, 71, 87, 92], 'attendance_question_grades': [0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0], 'final_PA_grade': 64.6, 'final_quiz_grade': 89.5, 'final_attendance_grade': 1.0, 'final_grade': 65.71, 'letter_grade': 'NP'}, '10012': {'PA_gradescope_scores': [80, 82, 88, 74, 93], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [80, 99, 72, 70, 86, 75, 73], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 0, 1, 1, 0, 0, 1, 1, 1, 1], 'final_PA_grade': 83.4, 'final_quiz_grade': 80.83333333333333, 'final_attendance_grade': 1.0, 'final_grade': 74.38999999999999, 'letter_grade': 'P'}, '10013': {'PA_gradescope_scores': [74, 97, 86, 90, 97], 'PA_late_days': [0, 9, 0, 0, 7], 'quiz_grades': [78, 96, 80, 77, 74, 99, 76], 'attendance_question_grades': [1, 1, 0, 1, 1, 1, 1, 1, 1, 0, 1, 1, 0, 1, 1], 'final_PA_grade': 50.0, 'final_quiz_grade': 84.33333333333333, 'final_attendance_grade': 1.0, 'final_grade': 55.4, 'letter_grade': 'NP'}, '10014': {'PA_gradescope_scores': [92, 76, 91, 100, 79], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [75, 70, 78, 89, 76, 93, 83], 'attendance_question_grades': [0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 1, 0], 'final_PA_grade': 87.6, 'final_quiz_grade': 82.33333333333333, 'final_attendance_grade': 1.0, 'final_grade': 77.35999999999999, 'letter_grade': 'P'}, '10015': {'PA_gradescope_scores': [87, 100, 99, 96, 78], 'PA_late_days': [0, 7, 0, 0, 0], 'quiz_grades': [83, 73, 74, 83, 97, 98, 99], 'attendance_question_grades': [1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1], 'final_PA_grade': 72.0, 'final_quiz_grade': 89.0, 'final_attendance_grade': 1.0, 'final_grade': 69.99999999999999, 'letter_grade': 'NP'}, '10016': {'PA_gradescope_scores': [93, 82, 84, 73, 73], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [92, 81, 97, 97, 90, 98, 75], 'attendance_question_grades': [1, 1, 1, 0, 1, 0, 1, 1, 1, 1, 0, 1, 1, 1, 1], 'final_PA_grade': 81.0, 'final_quiz_grade': 92.5, 'final_attendance_grade': 1.0, 'final_grade': 76.44999999999999, 'letter_grade': 'P'}, '10017': {'PA_gradescope_scores': [89, 99, 82, 74, 82], 'PA_late_days': [0, 6, 9, 6, 0], 'quiz_grades': [84, 76, 71, 77, 83, 70, 100], 'attendance_question_grades': [1, 0, 1, 0, 0, 1, 1, 1, 1, 1, 0, 1, 1, 0, 1], 'final_PA_grade': 34.2, 'final_quiz_grade': 81.83333333333333, 'final_attendance_grade': 1.0, 'final_grade': 45.169999999999995, 'letter_grade': 'NP'}, '10018': {'PA_gradescope_scores': [71, 94, 93, 80, 86], 'PA_late_days': [0, 0, 0, 0, 0], 'quiz_grades': [97, 80, 94, 95, 78, 73, 95], 'attendance_question_grades': [1, 1, 1, 0, 1, 0, 1, 1, 1, 0, 1, 1, 1, 1, 0], 'final_PA_grade': 84.8, 'final_quiz_grade': 89.83333333333333, 'final_attendance_grade': 1.0, 'final_grade': 77.92999999999999, 'letter_grade': 'P'}, '10019': {'PA_gradescope_scores': [78, 77, 98, 88, 84], 'PA_late_days': [7, 0, 0, 0, 0], 'quiz_grades': [80, 74, 80, 90, 74, 70, 83], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 0, 1, 1, 1, 0, 0, 0, 0, 1], 'final_PA_grade': 69.4, 'final_quiz_grade': 80.16666666666667, 'final_attendance_grade': 1.0, 'final_grade': 65.78999999999999, 'letter_grade': 'NP'}, '10020': {'PA_gradescope_scores': [83, 87, 76, 84, 71], 'PA_late_days': [0, 0, 10, 0, 0], 'quiz_grades': [78, 71, 81, 95, 91, 99, 81], 'attendance_question_grades': [1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1], 'final_PA_grade': 65.0, 'final_quiz_grade': 87.5, 'final_attendance_grade': 1.0, 'final_grade': 65.35, 'letter_grade': 'NP'}}
```
