```python
# Initialize Otter
import otter
grader = otter.Notebook("pa4.ipynb")
```

```python
import csv
import json
import sys
import io
import random
import math

# Utility function for testing
# You can ignore
def capture_print_output(func, *args, **kwargs):
    buffer = io.StringIO()
    old_stdout = sys.stdout
    sys.stdout = buffer
    try:
        func(*args, **kwargs)
    finally:
        sys.stdout = old_stdout
    return buffer.getvalue()
```

# Programming Assignment 4 Binary Search Trees and Smartphones Data Cleaning

In this programming assignment, you’ll work with a smartphone dataset in JSON format that needs cleaning and organizing. You’ll build a Binary Search Tree (BST) to efficiently store and retrieve smartphone data. The assignment has three main parts: (1) building a BST that can store both a key value and associated data within each node; (2) loading and cleaning JSON smartphone data to prepare it for insertion into the BST; and (3) using the BST to search for and retrieve phones based on their prices and other attributes. This project will help you develop skills in data cleaning, JSON handling, tree structures, and efficient data retrieval techniques.

## Part 1: Binary Search Tree (BST) Implementation

The `BSTNode` class includes a `key` attribute (equivalent to the `value` attribute in the lecture notes) to serve as each node’s unique, sortable identifier. Additionally, the class contains a `data` attribute, which is a list for holding multiple entries associated with the `key`. In the context of the smartphone data, each BST node uses the price as its `key` and stores a list of phone models of the same price in the `data` attribute. You can follow the BST implementation covered in the lecture. There is no need to balance the tree.

### Task 1: Implement the `__init__` Method

Initializes a `BSTNode` with a given `key` and optional initial `data`. If data is provided, it will be added to the node’s `data` list; otherwise, an empty list is created.

Parameters:
* `key`: The key for the node, used to order nodes in the BST.
* `data` (optional): The initial data to associate with the key. If provided, it is added to the `data` list; if not, an empty list is created.

```python
class BSTNode: 
    ...
```

```python
# Examples: 
price100 = BSTNode(100)
print(price100.key)  # 100
print(price100.data)  # []

price150 = BSTNode(150, {"model": "ModelA"})
print(price150.key)  # 150
print(price150.data)  # [{'model': 'ModelA'}]
```

```python
grader.check("P1T1")
```

### Task 2: Implement the `insert` Method

Adds a new key and associated data into the BST. If the key already exists, appends the data to that node’s data list.

Parameters:
* `new_key`: The key to insert. Determines the position of the new node in the BST.
* `associated_data`: The data to associate with `new_key`. This is appended to the `data` list of the existing node if the key already exists.

```python
# Examples: 
root = BSTNode(100, {"model": "RootModel"})

root.insert(50, {"model": "ModelA"})   
root.insert(150, {"model": "ModelB"})   
root.insert(25, {"model": "ModelC"})  
root.insert(75, {"model": "ModelD"})   
root.insert(75, {"model": "ModelE"}) 

print(root.key)               # 100
print(root.data)              # [{'model': 'RootModel'}]
print(root.left.key)          # 50
print(root.left.data)         # [{'model': 'ModelA'}]
print(root.left.left.key)     # 25
print(root.left.left.data)    # [{'model': 'ModelC'}]
print(root.left.right.key)    # 75
print(root.left.right.data)   # [{'model': 'ModelD'}, {'model': 'ModelE'}]
print(root.right.key)         # 150
print(root.right.data)        # [{'model': 'ModelB'}]
```

```python
grader.check("P1T2")
```

### Task 3: Implement the `print` Method

The `print` method performs an in-order traversal of the BST, printing each node's key and associated data on a separate line. 

```python
# Examples: 
root.print()
# Expected results: 
# 25: [{'model': 'ModelC'}]
# 50: [{'model': 'ModelA'}]
# 75: [{'model': 'ModelD'}, {'model': 'ModelE'}]
# 100: [{'model': 'RootModel'}]
# 150: [{'model': 'ModelB'}]
```

```python
grader.check("P1T3")
```

### Task 4: Implement the `search` Method

Searches the BST for a node with the specified key and returns the associated data list if found. If the key does not exist in the BST, the method returns None.

Parameters:
* `target_key`: The key to search for in the BST.

Returns:
* The list of associated data for the node with the specified key if it exists; otherwise, `None`. 

```python
# Examples: 
print(root.search(100))  # [{'model': 'RootModel'}]
print(root.search(50))   # [{'model': 'ModelA'}]
print(root.search(150))  # [{'model': 'ModelB'}]
print(root.search(25))   # [{'model': 'ModelC'}]
print(root.search(75))   # [{'model': 'ModelD'}, {'model': 'ModelE'}]
print(root.search(200))  # None
```

```python
grader.check("P1T4")
```

### Task 5: Implement the `count_entries` Method

Returns the total number of entries in the entire BST, including all entries in the `data` attribute of every node. 

```python
# Examples: 
root.count_entries()  # 6
# Explanation:
#   There are six entries: 
#   {"model": "RootModel"}, {"model": "ModelA"}, {"model": "ModelB"}, {"model": "ModelC"}, {"model": "ModelD"}, and {"model": "ModelE"}
```

```python
grader.check("P1T5")
```

### Task 6: Implement the `find_min` and `find_max` Methods

Find the node with the smallest or largest key in the BST. 

Returns:
* A tuple containing the key and its associated data for the minimum or maximum node.

```python
# Examples: 
print(root.find_min())  # (25, [{'model': 'ModelC'}])
print(root.find_max())  # (150, [{'model': 'ModelB'}])
```

```python
grader.check("P1T6")
```

### Task 7: Implement the `find_range` Method

Finds all nodes in the BST whose keys fall within a specified range. The method returns a list of tuples, each containing the key and its associated data for nodes in the range.

Parameters:
* `low`: The lower bound of the range. Inclusive.
* `high`: The upper bound of the range. **Exclusive**.

Returns:
* A list of tuples, where each tuple contains the key and its associated data for nodes within the range.

```python
# Examples: 
root.find_range(25, 75)
# Expected results: 
# [(25, [{'model': 'ModelC'}]), (50, [{'model': 'ModelA'}])]
```

```python
grader.check("P1T7")
```

### Task 8: Implement `build_bst_from_list` Function

This function takes a list of dictionaries, where each dictionary contains information about a phone (e.g., price and model), and builds a Binary Search Tree (BST) based on the phone prices. 

Parameters: 
* `phone_list`: A list of dictionaries, where each dictionary contains information about a phone, including a `price` key and other details.

Returns:
* The root node of the BST created from the phone_list.

```python
def build_bst_from_list(phone_list):
    ...
```

```python
# Examples:
phones = [
    {"price": 100, "model": "RootModel"},
    {"price": 50, "model": "ModelA"},
    {"price": 150, "model": "ModelB"},
    {"price": 25, "model": "ModelC"},
    {"price": 75, "model": "ModelD"},
    {"price": 75, "model": "ModelE"},
]

root = build_bst_from_list(phones)
root.print()
# Expected results:
# 25: [{'price': 25, 'model': 'ModelC'}]
# 50: [{'price': 50, 'model': 'ModelA'}]
# 75: [{'price': 75, 'model': 'ModelD'}, {'price': 75, 'model': 'ModelE'}]
# 100: [{'price': 100, 'model': 'RootModel'}]
# 150: [{'price': 150, 'model': 'ModelB'}]
```

```python
grader.check("P1T8")
```

## Part 2: Smartphone Data Cleaning

The dataset is downloaded from [here](https://www.kaggle.com/datasets/jakubkhalponiak/phones-2024) under [Database Contents License](https://opendatacommons.org/licenses/dbcl/1-0/). The data was scraped from GSMArena.com, including information on various phones currently available on the market. 

### Task 1: Load data
1. Load `data.json` and save it in `phone_data`.
   **NOTE:** When opening the file, add the parameter `encoding='utf-8'` to address issues with Unicode characters, such as currency symbols in the price field, which we will see later. When opening files, specifying an encoding tells Python how to interpret characters in the file. UTF-8 is a common encoding that supports a wide range of characters, including symbols like €, £, and ¥, by using 1 to 4 bytes per character. Unlike ASCII, which only supports basic English characters, UTF-8 can handle text from most languages and is ideal for files with diverse symbols.
3. Save its 8th (0-indexed) element in `phone8`.
4. Save the total number of records in `num_record`.

```python
...

print(num_record)  # 8791
phone8
# Expected results: 
# {'phone_brand': 'oukitel',
#  'phone_model': 'Oukitel RT2',
#  'price': 'About 350 EUR',
#  'specs': {'Network': {'2G bands': 'GSM 850 / 900 / 1800 / 1900 - SIM 1 & SIM 2',
#    '3G bands': 'HSDPA 850 / 900 / 1700(AWS) / 1900 / 2100 ',
#    '4G bands': ' LTE',
#    'Speed': 'HSPA, LTE'},
#   'Launch': {'Announced': '2022, September',
#    'Status': 'Available. Released 2022, September'},
#   'Body': {'Dimensions': '-',
#    'Weight': '-',
#    'SIM': 'Hybrid Dual SIM (Nano-SIM, dual stand-by)'},
#   'Display': {'Type': 'IPS LCD, 350 nits',
#    'Size': '10.1 inches, 295.8 cm',
#    'Resolution': '1200 x 1920 pixels, 16:10 ratio (~224 ppi density)'},
#   'Platform': {'OS': 'Android 12',
#    'Chipset': 'Mediatek MT8788 (12 nm)',
#    'CPU': 'Octa-core (4x2.0 GHz Cortex-A73 & 4x2.0 GHz Cortex-A53)',
#    'GPU': 'Mali-G72 MP3'},
#   'Memory': {'Card slot': 'microSDXC (uses shared SIM slot)',
#    'Internal': '128GB 8GB RAM'},
#   'Main Camera': {'Single': '16 MP',
#    'Features': 'LED flash',
#    'Video': '1080p@30fps'},
#   'Selfie camera': {'Single': '16 MP', 'Video': 'Yes'},
#   'Sound': {'Loudspeaker': 'Yes, with stereo speakers',
#    '3.5mm jack': 'Unspecified'},
#   'Comms': {'WLAN': 'Wi-Fi 802.11 a/b/g/n/ac, dual-band',
#    'Bluetooth': '4.2, A2DP',
#    'Positioning': 'GPS, GLONASS, GALILEO',
#    'NFC': 'No',
#    'Radio': 'Unspecified',
#    'USB': 'USB Type-C 2.0, OTG'},
#   'Features': {'Sensors': 'Unspecified'},
#   'Battery': {'Type': 'Li-Po 20000 mAh, non-removable',
#    'Charging': '33W wired'},
#   'Misc': {'Colors': 'Black, Orange', 'Price': 'About 350 EUR'}}}
```

```python
grader.check("P2T1")
```

### Task 2: Count the number of phone models for each phone brand
1. Count the number of records for each phone brand and sort the results in descending order.
2. Save the results in `phone_brand_counts`. 

```python
...
phone_brand_counts
# Expected results: 
# {'samsung': 1400,
#  'lg': 667,
#  'motorola': 634,
#  'nokia': 585,
#  'vivo': 476,
#  'huawei': 460,
#  'xiaomi': 412,
#  'alcatel': 409,
#  'zte': 387,
#  'oppo': 318,
#  'micromax': 289,
#  'htc': 289,
#  'lenovo': 249,
#  'honor': 227,
#  'realme': 221,
#  'asus': 200,
#  'sony': 160,
#  'tecno': 153,
#  'infinix': 140,
#  'apple': 128,
#  'ulefone_': 99,
#  'doogee': 92,
#  'blackview': 80,
#  'meizu': 78,
#  'cubot': 77,
#  'energizer': 77,
#  'oneplus': 77,
#  'tcl': 75,
#  'sharp': 75,
#  'umidigi': 71,
#  'oukitel': 55,
#  'itel': 39,
#  'google': 33,
#  'microsoft': 32,
#  'cat': 22,
#  'nothing': 5}
```

```python
grader.check("P2T2")
```

### Task 3: Remove records with no prices and compute percenatge of removed entries by brand

1. Filter out records with `None` in the `price` field and save the remaining data to `phone_data_with_prices`.
2. Track the percentage of removals per brand in `removed_percenatge_by_brand`. For example, Samsung has 1400 different models and 245 of them don't have prices. Then the removal percentage is 245 / 1400. Hint: Use `phone_brand_counts` that you generated in the previous task. 
3. Sort `removed_percenatge_by_brand` by removal percentage.

When working with real-world data, missing values are a common challenge. They can be addressed in various ways, with common practices including removing the empty entries or replacing missing values with default values or averages. Sometimes, you may also want to analyze the missing data, such as identifying the brands for which prices are more likely to be unavailable.

```python
...
print(len(phone_data_with_prices)) # 7183
removed_percenatge_by_brand

# Expected results: 
# {'umidigi': 0.8309859154929577,
#  'cubot': 0.6753246753246753,
#  'itel': 0.6666666666666666,
#  'alcatel': 0.5281173594132029,
#  'sharp': 0.4666666666666667,
#  'tecno': 0.45751633986928103,
#  'energizer': 0.33766233766233766,
#  'infinix': 0.3357142857142857,
#  'ulefone_': 0.29292929292929293,
#  'zte': 0.26356589147286824,
#  'motorola': 0.24605678233438485,
#  'lg': 0.21889055472263869,
#  'nokia': 0.18974358974358974,
#  'tcl': 0.18666666666666668,
#  'samsung': 0.175,
#  'doogee': 0.16304347826086957,
#  'blackview': 0.1625,
#  'sony': 0.14375,
#  'micromax': 0.13494809688581316,
#  'asus': 0.11,
#  'oukitel': 0.10909090909090909,
#  'microsoft': 0.09375,
#  'htc': 0.09342560553633218,
#  'huawei': 0.07391304347826087,
#  'meizu': 0.0641025641025641,
#  'vivo': 0.05672268907563025,
#  'lenovo': 0.05622489959839357,
#  'oppo': 0.050314465408805034,
#  'honor': 0.039647577092511016,
#  'oneplus': 0.03896103896103896,
#  'xiaomi': 0.03155339805825243,
#  'realme': 0.02262443438914027}
```

```python
grader.check("P2T3")
```

### Task 4: Implement the `get_random_records_by_field` function

This function extracts a random subset of phone records and returns only a specific field (e.g., price) for each record in a list. 

Parameters:
* `phone_list`: A list of dictionaries representing phone records.
* `field`: The key (field) in the dictionary to display for each selected record.
* `count`: The number of random records to return.

Returns:
* A list of values corresponding to the specified field from each randomly selected record.

You must use `random.sample` for random selection.

```python
def get_random_records_by_field(phone_list, field, count):
    ...
```

```python
random.seed(105)
get_random_records_by_field(phone_data_with_prices, 'price', 50)

# Expected results: 
# ['€\u2009119.90 / $\u2009159.00 / £\u2009109.99',
#  'About 140 EUR',
#  '$\u2009107.00 / £\u2009130.00 / €\u2009134.99',
#  'About 130 EUR',
#  'About 140 EUR',
#  'About 280 EUR',
#  '$\u2009182.27 / €\u2009129.99 / £\u2009100.19 / ₹\u20098,990',
#  'About 170 EUR',
#  'About 110 EUR',
#  '₹\u200922,499',
#  '€\u2009296.25 / $\u2009209.99 / £\u2009226.69',
#  'About 520 EUR',
#  'About 70 EUR',
#  '€\u2009122.78 / £\u2009149.99 / ₹\u200912,950',
#  'About 300 EUR',
#  'About 70 EUR',
#  'About 40 EUR',
#  '€\u2009481.58',
#  'About 200 EUR',
#  'About 30 EUR',
#  'About 70 EUR',
#  'About 2700 EUR',
#  'About 120 EUR',
#  '$\u200964.57 / €\u2009100.00 / £\u2009105.00',
#  'About 220 EUR',
#  'About 430 EUR',
#  'About 80 EUR',
#  'About 240 EUR',
#  'About 220 EUR',
#  '€\u2009169.09 / $\u2009410.00 / £\u2009358.50',
#  'About 830 EUR',
#  'About 160 EUR',
#  'About 550 EUR',
#  'About 90 EUR',
#  'About 160 EUR',
#  'About 200 EUR',
#  '£\u2009110.00 / €\u2009133.99',
#  'About 340 EUR',
#  'About 200 EUR',
#  'About 1070 EUR',
#  'About 110 EUR',
#  'About 230 EUR',
#  'About 200 EUR',
#  'About 180 EUR',
#  'About 60 EUR',
#  '$\u2009349.99 / C$\u2009576.23 / £\u2009214.99 / €\u2009259.99 / ₹\u200920,999',
#  'About 450 EUR',
#  'About 140 EUR',
#  'About 90 EUR',
#  'About 400 EUR']
```

```python
grader.check("P2T4")
```

### Task 5: Uniform prices

In the previous task, you observed that there are two formats for prices:

Type 1: 
* "About xxx EUR"
* "About xxx INR"

Type 2: 
* "€\u2009xxx"
* "€\u2009xxx / $\u2009xxx / £\u2009xxx"

Here, "\u2009" represents the Unicode character for a thin space. For simplicity of the assignment, only these two formats would exist in the dataset. However, if you were working with real-world data, additional rounds of checking would be necessary.

After an exhaustive search, you’ll find the following currencies (values are corresponding ISO Currency Codes):
```python
{'EUR': 'EUR',
 'INR': 'INR',
 '$': 'USD',
 'A$': 'AUD',
 'C$': 'CAD',
 '£': 'GBP',
 '€': 'EUR',
 '₹': 'INR'}
```
 
Your task:

For the price field of each records in the `phone_list` argument: 
1. Extract the numerical prices and their corresponding currencies.
2. Convert the currencies to ISO Currency Codes using the dictionary provided above.
3. Generate a dictionary of prices in the form {currency_code: price} (this is important as Type 2 may have multiple prices).
4. Add this dictionary of prices using `price_dict` as the key to each records of `phone_dict`

```python
CURRENCY_MAP = {'EUR': 'EUR',
 'INR': 'INR',
 '$': 'USD',
 'A$': 'AUD',
 'C$': 'CAD',
 '£': 'GBP',
 '€': 'EUR',
 '₹': 'INR'
}

def process_prices(phone_list):
    ...

# Process prices for phone_data_with_prices
process_prices(phone_data_with_prices)
```

```python
# Examples: 
phone_list = [
    {"model": "1", "price": "About 300 EUR"},
    {"model": "2", "price": "About 1000 INR"},
    {"model": "3", "price": "$\u2009255.45 / C$\u2009340.20 / £\u2009254.99 / €\u2009470.89 / ₹\u200959,999"},
    {"model": "4", "price": "€\u2009481.58"}
]
process_prices(phone_list)
phone_list
# Expected results: 
# [{'model': '1', 'price': 'About 300 EUR', 'price_dict': {'EUR': 300.0}},
#  {'model': '2', 'price': 'About 1000 INR', 'price_dict': {'INR': 1000.0}},
#  {'model': '3',
#   'price': '$\u2009255.45 / C$\u2009340.20 / £\u2009254.99 / €\u2009470.89 / ₹\u200959,999',
#   'price_dict': {'USD': 255.45,
#    'CAD': 340.2,
#    'GBP': 254.99,
#    'EUR': 470.89,
#    'INR': 59999.0}},
#  {'model': '4', 'price': '€\u2009481.58', 'price_dict': {'EUR': 481.58}}]
```

```python
grader.check("P2T5")
```

### Task 6: Convert Prices to USD

Use exchange rates from `rate.json` (download from [here](https://www.floatrates.com/json-feeds.html), where `inverseRate` indicates the value of one unit of currency in USD) to convert the `price_dict` field in the phone list to USD. The converted values should be stored in a new field called `price_dict_usd`.

```python
def convert_prices_to_usd(phone_list):
    ...

# Covert prices for phone_data_with_prices
convert_prices_to_usd(phone_data_with_prices)
```

```python
# Examples: 
phone_list = [{'model': '1', 'price': 'About 300 EUR', 'price_dict': {'EUR': 300.0}},
 {'model': '2', 'price': 'About 1000 INR', 'price_dict': {'INR': 1000.0}},
 {'model': '3',
  'price': '$\u2009255.45 / C$\u2009340.20 / £\u2009254.99 / €\u2009470.89 / ₹\u200959,999',
  'price_dict': {'USD': 255.45,
   'CAD': 340.2,
   'GBP': 254.99,
   'EUR': 470.89,
   'INR': 59999.0}},
 {'model': '4', 'price': '€\u2009481.58', 'price_dict': {'EUR': 481.58}}]

convert_prices_to_usd(phone_list)
phone_list
# Expected results: 
# [{'model': '1',
#   'price': 'About 300 EUR',
#   'price_dict': {'EUR': 300.0},
#   'price_dict_usd': {'EUR': 316.90061281638}},
#  {'model': '2',
#   'price': 'About 1000 INR',
#   'price_dict': {'INR': 1000.0},
#   'price_dict_usd': {'INR': 11.845719567256}},
#  {'model': '3',
#   'price': '$\u2009255.45 / C$\u2009340.20 / £\u2009254.99 / €\u2009470.89 / ₹\u200959,999',
#   'price_dict': {'USD': 255.45,
#    'CAD': 340.2,
#    'GBP': 254.99,
#    'EUR': 470.89,
#    'INR': 59999.0},
#   'price_dict_usd': {'USD': 255.45,
#    'CAD': 242.11709275350654,
#    'GBP': 323.07422986001075,
#    'EUR': 497.41776523035054,
#    'INR': 710.7313283157928}},
#  {'model': '4',
#   'price': '€\u2009481.58',
#   'price_dict': {'EUR': 481.58},
#   'price_dict_usd': {'EUR': 508.70999040037424}}]
```

```python
grader.check("P2T6")
```

### Task 7: Build a BST Based on Prices

Return a Binary Search Tree (BST) using prices in USD from `phone_list`. For simplicity, if a phone has multiple prices, calculate and use the average price as the key for the BST. In real-life scenarios, handling multiple price listings would require a more in-depth analysis. Remember to use the `BSTNode` class you defined in Part 1. 

```python
def build_bst_from_usd_prices(phone_list):
    ...

# Build a BST for phone_data_with_prices
phone_bst = build_bst_from_usd_prices(phone_data_with_prices)
```

```python
# Examples: 

phone_list = [{'model': '1',
  'price': 'About 300 EUR',
  'price_dict': {'EUR': 300.0},
  'price_dict_usd': {'EUR': 316.90061281638}},
 {'model': '2',
  'price': 'About 1000 INR',
  'price_dict': {'INR': 1000.0},
  'price_dict_usd': {'INR': 11.845719567256}},
 {'model': '3',
  'price': '$\u2009255.45 / C$\u2009340.20 / £\u2009254.99 / €\u2009470.89 / ₹\u200959,999',
  'price_dict': {'USD': 255.45,
   'CAD': 340.2,
   'GBP': 254.99,
   'EUR': 470.89,
   'INR': 59999.0},
  'price_dict_usd': {'USD': 255.45,
   'CAD': 242.11709275350654,
   'GBP': 323.07422986001075,
   'EUR': 497.41776523035054,
   'INR': 710.7313283157928}},
 {'model': '4',
  'price': '€\u2009481.58',
  'price_dict': {'EUR': 481.58},
  'price_dict_usd': {'EUR': 508.70999040037424}}]

root = build_bst_from_usd_prices(phone_list)
root.print()
# Expected results: 
# 11.845719567256: [{'model': '2', 'price': 'About 1000 INR', 'price_dict': {'INR': 1000.0}, 'price_dict_usd': {'INR': 11.845719567256}}]
# 316.90061281638: [{'model': '1', 'price': 'About 300 EUR', 'price_dict': {'EUR': 300.0}, 'price_dict_usd': {'EUR': 316.90061281638}}]
# 405.7580832319321: [{'model': '3', 'price': '$\u2009255.45 / C$\u2009340.20 / £\u2009254.99 / €\u2009470.89 / ₹\u200959,999', 'price_dict': {'USD': 255.45, 'CAD': 340.2, 'GBP': 254.99, 'EUR': 470.89, 'INR': 59999.0}, 'price_dict_usd': {'USD': 255.45, 'CAD': 242.11709275350654, 'GBP': 323.07422986001075, 'EUR': 497.41776523035054, 'INR': 710.7313283157928}}]
# 508.70999040037424: [{'model': '4', 'price': '€\u2009481.58', 'price_dict': {'EUR': 481.58}, 'price_dict_usd': {'EUR': 508.70999040037424}}]
```

```python
grader.check("P2T7")
```

### Task 8. Find min and max in `phone_bst`

What's the min / max price in the `phone_bst` and what are the phone models with this price? Remeber to use corresponding methods in the `BSTNode` class. 

```python
...

print(min_price)  # 10.563353760545999
print(min_models) 
# ['Micromax X100', 'Micromax X101', 'Micromax X099', 'Micromax X098', 'Energizer Energy E10', 'Energizer Energy E10+', 
# 'Energizer Energy E12', 'Energizer E2', 'Energizer E3', 'Energizer E4', 'Energizer E13', 'Nokia 103', 'Nokia 105 (2019)']
print(max_price)  # 13732.3598887098
print(max_models)  # ['Apple Watch Edition 42mm (1st gen)']
```

```python
grader.check("P2T8")
```

### Task 9: Find phone models with the price between $1000-1050. 

Use the `phone_bst` and corresponding method and return a list of phone models. 

```python
...
phone_1000_1050
# Expected results: 
['Samsung Galaxy Tab S8+',
 'Apple iPad Air 13 (2024)',
 'Samsung Galaxy S24 Ultra',
 'Asus ROG Phone 7',
 'vivo NEX Dual Display',
 'ZTE nubia Red Magic 6 Pro',
 'Huawei Pocket 2',
 'Samsung Galaxy Z Fold2 5G',
 'Honor Magic Vs3',
 'Asus ROG Phone 8',
 'Motorola Razr 50 Ultra']
```

```python
grader.check("P2T9")
```

### Task 10: Count phone models by price bracket

Determine the number of phone models within specific price brackets using the `phone_bst`. The price brackets are 0 to 2000 (exclusive) with a step of 100. That is (inclusive) 1-100 (exclusive), (inclusive) 100-200 (exclusive), 200-300, ..., 1900-2000. 

```python
...
brackets_count
# Expect results: 
# {(0, 100): 1486,
#  (100, 200): 2343,
#  (200, 300): 1606,
#  (300, 400): 714,
#  (400, 500): 409,
#  (500, 600): 201,
#  (600, 700): 113,
#  (700, 800): 82,
#  (800, 900): 62,
#  (900, 1000): 39,
#  (1000, 1100): 33,
#  (1100, 1200): 21,
#  (1200, 1300): 14,
#  (1300, 1400): 12,
#  (1400, 1500): 8,
#  (1500, 1600): 10,
#  (1600, 1700): 5,
#  (1700, 1800): 2,
#  (1800, 1900): 3,
#  (1900, 2000): 1}
```

```python
grader.check("P2T10")
```

## Conclusion

In this programming assignment, you’ve built a Binary Search Tree (BST) to organize and analyze phone data and explored cleaning and preparing the dataset for analysis. In the next assignment, we’ll use CSV files and the Pandas library to further clean and organize the dataset and create visualizations to gain deeper insights into the data.

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for P1T1 (points: 1)
```
>>> assert price100.key == 100
>>> assert price100.data == []
>>> assert price150.key == 150
>>> assert price150.data == [{'model': 'ModelA'}]
```


Tests for P1T2 (points: 1)
```
>>> assert root.key == 100
>>> assert root.data == [{'model': 'RootModel'}]
>>> assert root.left.key == 50
>>> assert root.left.data == [{'model': 'ModelA'}]
>>> assert root.left.left.key == 25
>>> assert root.left.left.data == [{'model': 'ModelC'}]
>>> assert root.left.right.key == 75
>>> assert root.left.right.data == [{'model': 'ModelD'}, {'model': 'ModelE'}]
>>> assert root.right.key == 150
>>> assert root.right.data == [{'model': 'ModelB'}]
```


Tests for P1T3 (points: 1)
```
>>> output = capture_print_output(root.print)
>>> assert output == "25: [{'model': 'ModelC'}]\n50: [{'model': 'ModelA'}]\n75: [{'model': 'ModelD'}, {'model': 'ModelE'}]\n100: [{'model': 'RootModel'}]\n150: [{'model': 'ModelB'}]\n"
```


Tests for P1T4 (points: 1)
```
>>> assert root.search(100) == [{'model': 'RootModel'}]
>>> assert root.search(50) == [{'model': 'ModelA'}]
>>> assert root.search(150) == [{'model': 'ModelB'}]
>>> assert root.search(25) == [{'model': 'ModelC'}]
>>> assert root.search(75) == [{'model': 'ModelD'}, {'model': 'ModelE'}]
>>> assert root.search(200) is None
```


Tests for P1T5 (points: 1)
```
>>> assert root.count_entries() == 6
```


Tests for P1T6 (points: 1)
```
>>> assert root.find_min() == (25, [{'model': 'ModelC'}])
>>> assert root.find_max() == (150, [{'model': 'ModelB'}])
```


Tests for P1T7 (points: 1)
```
>>> assert root.find_range(25, 75) == [(25, [{'model': 'ModelC'}]), (50, [{'model': 'ModelA'}])]
>>> assert root.find_range(25, 50) == [(25, [{'model': 'ModelC'}])]
>>> assert root.find_range(0, 200) == [(25, [{'model': 'ModelC'}]), (50, [{'model': 'ModelA'}]), (75, [{'model': 'ModelD'}, {'model': 'ModelE'}]), (100, [{'model': 'RootModel'}]), (150, [{'model': 'ModelB'}])]
```


Tests for P1T8 (points: 1)
```
>>> phones_empty = []
>>> root_empty = build_bst_from_list(phones_empty)
>>> assert root_empty is None
>>> random.seed(42)
>>> phones_large = [{'price': random.randint(1, 300), 'model': f'Model_{i}'} for i in range(1, 201)]
>>> root_large = build_bst_from_list(phones_large)
>>> assert root_large is not None
>>> assert root_large.count_entries() == 200
>>> assert root_large.find_min() == (1, [{'price': 1, 'model': 'Model_145'}])
>>> assert root_large.find_max() == (300, [{'price': 300, 'model': 'Model_99'}])
>>> assert root_large.find_range(90, 100) == [(92, [{'price': 92, 'model': 'Model_134'}]), (99, [{'price': 99, 'model': 'Model_50'}])]
>>> assert root_large.find_range(200, 210) == [(203, [{'price': 203, 'model': 'Model_90'}]), (205, [{'price': 205, 'model': 'Model_102'}, {'price': 205, 'model': 'Model_173'}]), (206, [{'price': 206, 'model': 'Model_82'}])]
>>> def is_bst(root, min_value=float('-inf'), max_value=float('inf')):
...     if not root:
...         return True
...     if not min_value < root.key < max_value:
...         return False
...     left_is_bst = is_bst(root.left, min_value, root.key) if root.left else True
...     right_is_bst = is_bst(root.right, root.key, max_value) if root.right else True
...     return left_is_bst and right_is_bst
>>> assert is_bst(root_large)
```


Tests for P2T1 (points: 1)
```
>>> assert phone8 == {'phone_brand': 'oukitel', 'phone_model': 'Oukitel RT2', 'price': 'About 350 EUR', 'specs': {'Network': {'2G bands': 'GSM 850 / 900 / 1800 / 1900 - SIM 1 & SIM 2', '3G bands': 'HSDPA 850 / 900 / 1700(AWS) / 1900 / 2100 ', '4G bands': ' LTE', 'Speed': 'HSPA, LTE'}, 'Launch': {'Announced': '2022, September', 'Status': 'Available. Released 2022, September'}, 'Body': {'Dimensions': '-', 'Weight': '-', 'SIM': 'Hybrid Dual SIM (Nano-SIM, dual stand-by)'}, 'Display': {'Type': 'IPS LCD, 350 nits', 'Size': '10.1 inches, 295.8 cm', 'Resolution': '1200 x 1920 pixels, 16:10 ratio (~224 ppi density)'}, 'Platform': {'OS': 'Android 12', 'Chipset': 'Mediatek MT8788 (12 nm)', 'CPU': 'Octa-core (4x2.0 GHz Cortex-A73 & 4x2.0 GHz Cortex-A53)', 'GPU': 'Mali-G72 MP3'}, 'Memory': {'Card slot': 'microSDXC (uses shared SIM slot)', 'Internal': '128GB 8GB RAM'}, 'Main Camera': {'Single': '16 MP', 'Features': 'LED flash', 'Video': '1080p@30fps'}, 'Selfie camera': {'Single': '16 MP', 'Video': 'Yes'}, 'Sound': {'Loudspeaker': 'Yes, with stereo speakers', '3.5mm jack': 'Unspecified'}, 'Comms': {'WLAN': 'Wi-Fi 802.11 a/b/g/n/ac, dual-band', 'Bluetooth': '4.2, A2DP', 'Positioning': 'GPS, GLONASS, GALILEO', 'NFC': 'No', 'Radio': 'Unspecified', 'USB': 'USB Type-C 2.0, OTG'}, 'Features': {'Sensors': 'Unspecified'}, 'Battery': {'Type': 'Li-Po 20000 mAh, non-removable', 'Charging': '33W wired'}, 'Misc': {'Colors': 'Black, Orange', 'Price': 'About 350 EUR'}}}
>>> assert num_record == 8791
```


Tests for P2T10 (points: 1)
```
>>> assert brackets_count == {(0, 100): 1486, (100, 200): 2343, (200, 300): 1606, (300, 400): 714, (400, 500): 409, (500, 600): 201, (600, 700): 113, (700, 800): 82, (800, 900): 62, (900, 1000): 39, (1000, 1100): 33, (1100, 1200): 21, (1200, 1300): 14, (1300, 1400): 12, (1400, 1500): 8, (1500, 1600): 10, (1600, 1700): 5, (1700, 1800): 2, (1800, 1900): 3, (1900, 2000): 1}
```


Tests for P2T2 (points: 1)
```
>>> assert phone_brand_counts == {'samsung': 1400, 'lg': 667, 'motorola': 634, 'nokia': 585, 'vivo': 476, 'huawei': 460, 'xiaomi': 412, 'alcatel': 409, 'zte': 387, 'oppo': 318, 'micromax': 289, 'htc': 289, 'lenovo': 249, 'honor': 227, 'realme': 221, 'asus': 200, 'sony': 160, 'tecno': 153, 'infinix': 140, 'apple': 128, 'ulefone_': 99, 'doogee': 92, 'blackview': 80, 'meizu': 78, 'cubot': 77, 'energizer': 77, 'oneplus': 77, 'tcl': 75, 'sharp': 75, 'umidigi': 71, 'oukitel': 55, 'itel': 39, 'google': 33, 'microsoft': 32, 'cat': 22, 'nothing': 5}
```


Tests for P2T3 (points: 1)
```
>>> assert len(phone_data_with_prices) == 7183
>>> removed_percenatge_by_brand_expected = {'umidigi': 0.8309859154929577, 'cubot': 0.6753246753246753, 'itel': 0.6666666666666666, 'alcatel': 0.5281173594132029, 'sharp': 0.4666666666666667, 'tecno': 0.45751633986928103, 'energizer': 0.33766233766233766, 'infinix': 0.3357142857142857, 'ulefone_': 0.29292929292929293, 'zte': 0.26356589147286824, 'motorola': 0.24605678233438485, 'lg': 0.21889055472263869, 'nokia': 0.18974358974358974, 'tcl': 0.18666666666666668, 'samsung': 0.175, 'doogee': 0.16304347826086957, 'blackview': 0.1625, 'sony': 0.14375, 'micromax': 0.13494809688581316, 'asus': 0.11, 'oukitel': 0.10909090909090909, 'microsoft': 0.09375, 'htc': 0.09342560553633218, 'huawei': 0.07391304347826087, 'meizu': 0.0641025641025641, 'vivo': 0.05672268907563025, 'lenovo': 0.05622489959839357, 'oppo': 0.050314465408805034, 'honor': 0.039647577092511016, 'oneplus': 0.03896103896103896, 'xiaomi': 0.03155339805825243, 'realme': 0.02262443438914027}
>>> for brand, remove_percentage in removed_percenatge_by_brand.items():
...     assert math.isclose(remove_percentage, removed_percenatge_by_brand[brand])
```


Tests for P2T4 (points: 1)
```
>>> random.seed(368)
>>> assert get_random_records_by_field(phone_data_with_prices, 'price', 20) == ['About 150 EUR', 'About 150 EUR', 'About 20 EUR', 'About 30 EUR', 'About 60 EUR', 'About 20 EUR', 'About 300 EUR', 'About 300 EUR', 'About 300 EUR', 'About 300 EUR', 'About 180 EUR', 'About 450 EUR', 'About 390 EUR', 'About 250 EUR', '$\u2009718.00 / €\u2009933.97', '€\u2009422.74 / $\u2009332.00 / £\u2009305.00', 'About 160 EUR', 'About 110 EUR', 'About 150 EUR', 'About 170 EUR']
>>> random.seed(369)
>>> assert get_random_records_by_field(phone_data_with_prices, 'price', 20) == ['About 130 EUR', 'About 90 EUR', 'About 90 EUR', '€\u2009334.78 / £\u2009189.00 / ₹\u200916,999', 'About 430 EUR', 'About 210 EUR', 'About 330 EUR', '$\u2009146.43 / £\u2009177.11 / €\u2009196.00', 'About 200 EUR', 'About 140 EUR', 'About 960 EUR', 'About 170 EUR', 'About 600 EUR', '$\u2009255.45 / C$\u2009340.20 / £\u2009254.99 / €\u2009470.89 / ₹\u200959,999', 'About 500 EUR', 'About 20 EUR', 'About 260 EUR', 'About 150 EUR', '€\u2009139.99 / $\u2009160.00 / £\u2009115.99', 'About 350 EUR']
```


Tests for P2T5 (points: 1)
```
>>> assert phone_list == [{'model': '1', 'price': 'About 300 EUR', 'price_dict': {'EUR': 300.0}}, {'model': '2', 'price': 'About 1000 INR', 'price_dict': {'INR': 1000.0}}, {'model': '3', 'price': '$\u2009255.45 / C$\u2009340.20 / £\u2009254.99 / €\u2009470.89 / ₹\u200959,999', 'price_dict': {'USD': 255.45, 'CAD': 340.2, 'GBP': 254.99, 'EUR': 470.89, 'INR': 59999.0}}, {'model': '4', 'price': '€\u2009481.58', 'price_dict': {'EUR': 481.58}}]
>>> random.seed(300)
>>> assert [{'phone_model': r['phone_model'], 'price_dict': r['price_dict']} for r in random.sample(phone_data_with_prices, 10)] == [{'phone_model': 'Huawei MediaPad S7-301w', 'price_dict': {'EUR': 160.0}}, {'phone_model': 'OnePlus Nord', 'price_dict': {'USD': 199.99, 'GBP': 123.0, 'EUR': 159.99}}, {'phone_model': 'Oppo F25 Pro', 'price_dict': {'EUR': 270.0}}, {'phone_model': 'Doogee N30', 'price_dict': {'EUR': 150.0}}, {'phone_model': 'Lenovo Tab S8', 'price_dict': {'EUR': 150.0}}, {'phone_model': 'Oppo A3x (China)', 'price_dict': {'EUR': 150.0}}, {'phone_model': 'HTC U Ultra', 'price_dict': {'EUR': 250.0}}, {'phone_model': 'Samsung Galaxy Note5 Duos', 'price_dict': {'EUR': 450.0}}, {'phone_model': 'LG Optimus LTE2', 'price_dict': {'EUR': 210.0}}, {'phone_model': 'Xiaomi 11T Pro', 'price_dict': {'GBP': 159.99, 'EUR': 182.99, 'INR': 19499.0}}]
>>> random.seed(800)
>>> assert [{'phone_model': r['phone_model'], 'price_dict': r['price_dict']} for r in random.sample(phone_data_with_prices, 10)] == [{'phone_model': 'Nokia 6086', 'price_dict': {'EUR': 130.0}}, {'phone_model': 'Motorola Razr 2022', 'price_dict': {'USD': 940.0, 'EUR': 669.99, 'GBP': 749.0}}, {'phone_model': 'Xiaomi Poco M6 4G', 'price_dict': {'USD': 154.69, 'EUR': 125.0, 'GBP': 113.74, 'INR': 13490.0}}, {'phone_model': 'Micromax X101', 'price_dict': {'EUR': 10.0}}, {'phone_model': 'Lenovo IdeaTab A2107', 'price_dict': {'EUR': 150.0}}, {'phone_model': 'Motorola RAZR i XT890', 'price_dict': {'EUR': 180.0}}, {'phone_model': 'Samsung Galaxy 551', 'price_dict': {'EUR': 90.0}}, {'phone_model': 'HTC Desire 820 dual sim', 'price_dict': {'EUR': 320.0}}, {'phone_model': 'Oppo Find N2 Flip', 'price_dict': {'EUR': 699.99, 'GBP': 400.0}}, {'phone_model': 'Nokia 105 (2019)', 'price_dict': {'EUR': 10.0}}]
>>> random.seed(900)
>>> assert [{'phone_model': r['phone_model'], 'price_dict': r['price_dict']} for r in random.sample(phone_data_with_prices, 10)] == [{'phone_model': 'Huawei nova 4', 'price_dict': {'EUR': 400.0}}, {'phone_model': 'Meizu 15', 'price_dict': {'EUR': 400.0}}, {'phone_model': 'Motorola W220', 'price_dict': {'EUR': 50.0}}, {'phone_model': 'Samsung E1120', 'price_dict': {'EUR': 30.0}}, {'phone_model': 'Tecno Camon 30 Premier', 'price_dict': {'INR': 39999.0}}, {'phone_model': 'Nokia 2220 slide', 'price_dict': {'EUR': 50.0}}, {'phone_model': 'Infinix Hot 20i', 'price_dict': {'EUR': 120.0}}, {'phone_model': 'Nokia Asha 503 Dual SIM', 'price_dict': {'EUR': 100.0}}, {'phone_model': 'Asus Zenfone 2E', 'price_dict': {'EUR': 110.0}}, {'phone_model': 'Sony Xperia Pro-I', 'price_dict': {'EUR': 930.0}}]
>>> random.seed(1500)
>>> assert [{'phone_model': r['phone_model'], 'price_dict': r['price_dict']} for r in random.sample(phone_data_with_prices, 10)] == [{'phone_model': 'Oppo K10 5G', 'price_dict': {'EUR': 270.0}}, {'phone_model': 'Micromax A121 Canvas Elanza 2', 'price_dict': {'EUR': 120.0}}, {'phone_model': 'Samsung I8200 Galaxy S III mini VE', 'price_dict': {'EUR': 180.0}}, {'phone_model': 'HTC 7 Trophy', 'price_dict': {'EUR': 90.0}}, {'phone_model': 'Realme 12x', 'price_dict': {'USD': 162.05}}, {'phone_model': 'Infinix Note 40 Pro 4G', 'price_dict': {'EUR': 240.0}}, {'phone_model': 'Realme C55', 'price_dict': {'USD': 275.0, 'EUR': 109.27, 'INR': 7498.0}}, {'phone_model': 'LG KG810', 'price_dict': {'EUR': 120.0}}, {'phone_model': 'Samsung S5500 Eco', 'price_dict': {'EUR': 150.0}}, {'phone_model': 'Honor Play4 Pro', 'price_dict': {'EUR': 360.0}}]
```


Tests for P2T6 (points: 1)
```
>>> assert phone_list == [{'model': '1', 'price': 'About 300 EUR', 'price_dict': {'EUR': 300.0}, 'price_dict_usd': {'EUR': 316.90061281638}}, {'model': '2', 'price': 'About 1000 INR', 'price_dict': {'INR': 1000.0}, 'price_dict_usd': {'INR': 11.845719567256}}, {'model': '3', 'price': '$\u2009255.45 / C$\u2009340.20 / £\u2009254.99 / €\u2009470.89 / ₹\u200959,999', 'price_dict': {'USD': 255.45, 'CAD': 340.2, 'GBP': 254.99, 'EUR': 470.89, 'INR': 59999.0}, 'price_dict_usd': {'USD': 255.45, 'CAD': 242.11709275350654, 'GBP': 323.07422986001075, 'EUR': 497.41776523035054, 'INR': 710.7313283157928}}, {'model': '4', 'price': '€\u2009481.58', 'price_dict': {'EUR': 481.58}, 'price_dict_usd': {'EUR': 508.70999040037424}}]
>>> random.seed(300)
>>> assert [{'phone_model': r['phone_model'], 'price_dict_usd': r['price_dict_usd']} for r in random.sample(phone_data_with_prices, 10)] == [{'phone_model': 'Huawei MediaPad S7-301w', 'price_dict_usd': {'EUR': 169.01366016873598}}, {'phone_model': 'OnePlus Nord', 'price_dict_usd': {'USD': 199.99, 'GBP': 155.8419164390028, 'EUR': 169.00309681497546}}, {'phone_model': 'Oppo F25 Pro', 'price_dict_usd': {'EUR': 285.21055153474197}}, {'phone_model': 'Doogee N30', 'price_dict_usd': {'EUR': 158.45030640819}}, {'phone_model': 'Lenovo Tab S8', 'price_dict_usd': {'EUR': 158.45030640819}}, {'phone_model': 'Oppo A3x (China)', 'price_dict_usd': {'EUR': 158.45030640819}}, {'phone_model': 'HTC U Ultra', 'price_dict_usd': {'EUR': 264.08384401365}}, {'phone_model': 'Samsung Galaxy Note5 Duos', 'price_dict_usd': {'EUR': 475.35091922456996}}, {'phone_model': 'LG Optimus LTE2', 'price_dict_usd': {'EUR': 221.83042897146598}}, {'phone_model': 'Xiaomi 11T Pro', 'price_dict_usd': {'GBP': 202.70852204126876, 'EUR': 193.29881046423125, 'INR': 230.97968584192475}}]
>>> random.seed(800)
>>> assert [{'phone_model': r['phone_model'], 'price_dict_usd': r['price_dict_usd']} for r in random.sample(phone_data_with_prices, 10)] == [{'phone_model': 'Nokia 6086', 'price_dict_usd': {'EUR': 137.323598887098}}, {'phone_model': 'Motorola Razr 2022', 'price_dict_usd': {'USD': 940.0, 'EUR': 707.7341386028214, 'GBP': 948.9885805919764}}, {'phone_model': 'Xiaomi Poco M6 4G', 'price_dict_usd': {'USD': 154.69, 'EUR': 132.041922006825, 'GBP': 144.10942744530226, 'INR': 159.79875696228345}}, {'phone_model': 'Micromax X101', 'price_dict_usd': {'EUR': 10.563353760545999}}, {'phone_model': 'Lenovo IdeaTab A2107', 'price_dict_usd': {'EUR': 158.45030640819}}, {'phone_model': 'Motorola RAZR i XT890', 'price_dict_usd': {'EUR': 190.140367689828}}, {'phone_model': 'Samsung Galaxy 551', 'price_dict_usd': {'EUR': 95.070183844914}}, {'phone_model': 'HTC Desire 820 dual sim', 'price_dict_usd': {'EUR': 338.02732033747196}}, {'phone_model': 'Oppo Find N2 Flip', 'price_dict_usd': {'EUR': 739.4241998844594, 'GBP': 506.80298028944}}, {'phone_model': 'Nokia 105 (2019)', 'price_dict_usd': {'EUR': 10.563353760545999}}]
>>> random.seed(1500)
>>> assert [{'phone_model': r['phone_model'], 'price_dict_usd': r['price_dict_usd']} for r in random.sample(phone_data_with_prices, 10)] == [{'phone_model': 'Oppo K10 5G', 'price_dict_usd': {'EUR': 285.21055153474197}}, {'phone_model': 'Micromax A121 Canvas Elanza 2', 'price_dict_usd': {'EUR': 126.760245126552}}, {'phone_model': 'Samsung I8200 Galaxy S III mini VE', 'price_dict_usd': {'EUR': 190.140367689828}}, {'phone_model': 'HTC 7 Trophy', 'price_dict_usd': {'EUR': 95.070183844914}}, {'phone_model': 'Realme 12x', 'price_dict_usd': {'USD': 162.05}}, {'phone_model': 'Infinix Note 40 Pro 4G', 'price_dict_usd': {'EUR': 253.520490253104}}, {'phone_model': 'Realme C55', 'price_dict_usd': {'USD': 275.0, 'EUR': 115.42576654148613, 'INR': 88.81920531528549}}, {'phone_model': 'LG KG810', 'price_dict_usd': {'EUR': 126.760245126552}}, {'phone_model': 'Samsung S5500 Eco', 'price_dict_usd': {'EUR': 158.45030640819}}, {'phone_model': 'Honor Play4 Pro', 'price_dict_usd': {'EUR': 380.280735379656}}]
```


Tests for P2T7 (points: 1)
```
>>> assert capture_print_output(root.print) == "11.845719567256: [{'model': '2', 'price': 'About 1000 INR', 'price_dict': {'INR': 1000.0}, 'price_dict_usd': {'INR': 11.845719567256}}]\n316.90061281638: [{'model': '1', 'price': 'About 300 EUR', 'price_dict': {'EUR': 300.0}, 'price_dict_usd': {'EUR': 316.90061281638}}]\n405.7580832319321: [{'model': '3', 'price': '$\\u2009255.45 / C$\\u2009340.20 / £\\u2009254.99 / €\\u2009470.89 / ₹\\u200959,999', 'price_dict': {'USD': 255.45, 'CAD': 340.2, 'GBP': 254.99, 'EUR': 470.89, 'INR': 59999.0}, 'price_dict_usd': {'USD': 255.45, 'CAD': 242.11709275350654, 'GBP': 323.07422986001075, 'EUR': 497.41776523035054, 'INR': 710.7313283157928}}]\n508.70999040037424: [{'model': '4', 'price': '€\\u2009481.58', 'price_dict': {'EUR': 481.58}, 'price_dict_usd': {'EUR': 508.70999040037424}}]\n"
```


Tests for P2T8 (points: 1)
```
>>> assert math.isclose(min_price, 10.563353760545999)
>>> assert min_models == ['Micromax X100', 'Micromax X101', 'Micromax X099', 'Micromax X098', 'Energizer Energy E10', 'Energizer Energy E10+', 'Energizer Energy E12', 'Energizer E2', 'Energizer E3', 'Energizer E4', 'Energizer E13', 'Nokia 103', 'Nokia 105 (2019)']
>>> assert math.isclose(max_price, 13732.3598887098)
>>> assert max_models == ['Apple Watch Edition 42mm (1st gen)']
```


Tests for P2T9 (points: 1)
```
>>> assert phone_1000_1050 == ['Samsung Galaxy Tab S8+', 'Apple iPad Air 13 (2024)', 'Samsung Galaxy S24 Ultra', 'Asus ROG Phone 7', 'vivo NEX Dual Display', 'ZTE nubia Red Magic 6 Pro', 'Huawei Pocket 2', 'Samsung Galaxy Z Fold2 5G', 'Honor Magic Vs3', 'Asus ROG Phone 8', 'Motorola Razr 50 Ultra']
```
